Single-cell transcriptomics identifies altered neutrophil dynamics and accentuated T-cell cytotoxicity in tobacco flavored e-cigarette exposed mouse lungs

Author response:

The following is the authors’ response to the previous reviews.

Public Reviews:

Reviewer #1 (Public review):

The authors assess the impact of E-cigarette smoke exposure on mouse lungs using single cell RNA sequencing. Air was used as control and several flavors (fruit, menthol, tobacco) were tested. Differentially expressed genes (DEGs) were identified for each group and compared against the air control. Changes in gene expression in either myeloid or lymphoid cells were identified for each flavor and the results varied by sex. The scRNAseq dataset will be of interest to the lung immunity and e-cig research communities and some of the observed effects could be important. Unfortunately, the revision did not address the reviewers' main concerns about low replicate numbers and lack of validations. The study remains preliminary, and no solid conclusions could be drawn about the effects of E-cig exposure as a whole or any flavor-specific phenotypes.

Strengths:

The study is the first to use scRNAseq to systematically analyze the impact of e-cigarettes on the lung. The dataset will be of broad interest.

Weaknesses:

scRNAseq studies may have low replicate numbers due to the high cost of studies but at least 2 or 3 biological replicates for each experimental group is required to ensure rigor of the interpretation. This study had only N=1 per sex per group and some sex-dependent effects were observed. This could have been remedied by validating key observations from the study using traditional methods such as flow cytometry and qPCR, but the limited number of validation experiments did not support the conclusions of the scRNA seq analysis. An important control group (PG:VG) had extremely low cell numbers and was basically not useful. Statistical analysis is lacking in almost all figures. Overall, this is a preliminary study with some potentially interesting observations, but no solid conclusions can be made from the data presented.

The only new validation experiment is the immunofluorescent staining of neutrophils in Figure 4. The images are very low resolution and low quality and it is not clear which cells are neutrophils. S100A8 (calprotectin) is highly abundant in neutrophils but not strictly neutrophil-specific. It's hard to distinguish positive cells from autofluorescence in both Ly6g and S100a8 channels. No statistical analysis in the quantification.

We thank the reviewer for identifying the strengths of this study and pointing out the gaps in knowledge. Overall, our purpose to present this data is to provide the scRNA seq results as a resource to a wider community. We have used techniques like flow cytometry, multianalyte cytokine array and immunofluorescence to validate some of the results. We agree with the reviewer that we were unable to rightly point out the significance of our findings with the immunofluorescent stain in the previous edit. We have revised the manuscript and included the quantification for both Ly6G+ and S100A8+ cells in e-cig aerosol exposed and control lung tissues. Briefly, we identified a marked decrease in the staining for S100A8 (marker for neutrophil activation) in tobacco-flavored e-cig exposed mouse lungs as compared to controls. Upon considering the corroborating evidence from scRNA seq and flow cytometry with regards to increased neutrophil percentages in experimental group and lowered staining for active neutrophils using immunofluorescence, we speculate that exposure to e-cig (tobacco) aerosols may alter the neutrophil dynamics within the lungs. Also, co-immunofluorescence identified a more prominent co-localization of the two markers in control samples as compared to the treatment group which points towards some changes in the innate immune milieu within the lungs upon exposures. Future work is required to validate these speculations.

We have now discussed all the above-mentioned points in the Discussion section of the revised manuscript and toned down our conclusions regarding sex-dependent changes from scRNA seq data.

It is unclear what the meaning of Fig. 3A and B is, since these numbers only reflect the number of cells captured in the scRNAseq experiment and are not biologically meaningful. Flow cytometry quantification is presented as cell counts, but the percentage of cells from the CD45+ gate should be shown. No statistical analysis is shown, and flow cytometry results do not support the conclusions of scRNAseq data.

We thank the reviewer for this question. However, we would like to highlight that scRNA seq and flow cytometry may show similar trends but cannot be identical as one relies on cell surface markers (protein) for identification of cell types, while other is dependent on the transcriptomic signatures to identify the cell types. In our data, for the myeloid cells (alveolar macrophages and neutrophils), the scRNA and flow cytometry data match in trend. However, the trends do not match with respect to the lymphoid cells being studied (CD4 and CD8 T cells). The possible explanation for such a finding could be possible high gene dropout rates in scRNA seq, different analytical resolution for the two techniques and pooling of samples in our single cell workflow. We realize these shortcomings in our analyses and mention it clearly in the discussion as limitation of our work. It is important to note also that cell frequencies identified in scRNA seq just provide wide and indistinct indications which need to be further validated, which we tried to accomplish in our work to some degree. Our flow-based results clearly highlight the sex-specific variations in the immune cell percentages (something we could not have anticipated earlier). In future studies, we will include more replicates to tease out sex-based variations upon acute and chronic exposure to e-cig aerosols.

We have now replotted the graphs in Fig 3A and B and plotted the flow quantification as the percentage of total CD45+ cells. The gating strategy for the flow plots is also included as Figure S6 in the revised manuscript.

Reviewer #2 (Public review):

This study provides some interesting observations on how different flavour e-cigarettes can affect lung immunology; however, there are numerous flaws, including a low replicate number and a lack of effective validation methods, meaning findings may not be repeated. This is a revised article but several weaknesses remain related to the analysis and interpretation of the data.

Strengths:

The strength of the study is the successful scRNA-seq experiment which gives some preliminary data that can be used to create new hypotheses in this area.

Weaknesses:

Although some text weaknesses have been addressed since resubmission, other specific weaknesses remain: The major weakness is the n-number and analysis methods. Two biological n per group is not acceptable to base any solid conclusions. Any validatory data was too little (only cell % data) and not always supporting the findings (e.g. figure 3D does not match 3B/4A). Other examples include:

There aren't enough cells to justify analysis - only 300-1500 myeloid cells per group with not many of these being neutrophils or the apparent 'Ly6G- neutrophils'.

We thank the reviewer for the comment, but we disagree with the reviewer in terms of the justification of analyses. All the flavored e-cig aerosol groups were compared with air controls to deduce the outcomes in the current study. We already acknowledge low sample quality for PGVG group and have only included the comparisons with PGVG upon reviewer’s request which is open to interpretation by the reader.

By that measure, each treatment group (except PGVG group) has over 1000 cells with 24777 genes being analyzed for each cell type, which by the standards of single cell is sufficient. We understand that this strategy should not be used for detection of rare cell populations, which was neither the purpose of this manuscript nor was attempted. We conduct comparisons of broader cell types and mention more samples need to be added in the Discussion section of the revised manuscript.

As for the Ly6G neutrophil category, we don’t only base our results on scRNA analyses but also perform co-immunofluorescence and multi-analyte analyses and use evidence from previous literature to back our outcome. To avoid over-stating our results we have revamped the whole manuscript and ensured to tone down our results with relation to the presence of Ly6G- neutrophils. We do understand that more work is required in the future, but our work clearly shows the shift in neutrophil dynamics upon exposure which should be reported, in our opinion.

The dynamic range of RNA measurement using scRNAseq is known to be limited - how do we know whether genes are not expressed or just didn't hit detection? This links into the Ly6G negative neutrophil comments, but in general the lack of gene expression in this kind of data should be viewed with caution, especially with a low n number and few cells. The data in the entire paper is not strong enough to base any solid conclusion - it is not just the RNA-sequencing data.

We acknowledge this to be a valid point and have revamped the manuscript and toned down our conclusions. However, such limitations exist with any scRNA seq dataset and so must be interpreted accordingly by the readers. We do understand that due to the low cell counts and the limitations with scRNA seq we should not perform DESeq2 analyses for Ly6G+ versus Ly6G- neutrophil categories, which was never attempted at the first place. However, our results with co-immunofluorescence, multianalyte assay and scRNA expression analyses in myeloid cluster do point towards a shift in neutrophil activation which needs to be further investigated. Furthermore, Ly6G deficiency has been linked to immature neutrophils in many previous studies and is not an unlikely outcome that needs to be treated with immense skepticism.

We wish to make this dataset available as a resource to influence future research. We are aware of its limitations and have been transparent with regards to our experimental design, capture strategy, the quality of obtained results, and possible caveats to make it is open for discussion by the readers.

There is no data supporting the presence of Ly6G negative neutrophils. In the flow cytometry only Ly6G+ cells are shown with no evidence of Ly6G negative neutrophils (assuming equal CD11b expression). There is no new data to support this claim since resubmission and the New figures 4C and D actually show there are no Ly6G negative cells - the cells that the authors deem Ly6G negative are actually positive - but the red overlay of S100A8 is so strong it blocks out the green signal - looking to the Ly6G single stains (green only) you can see that the reported S100A8+Ly6G- cells all have Ly6G (with different staining intensities).

We thank the reviewer for this query and do understand the skepticism. We have now quantified the data to provide more clarity for interpretation. As we were using paraffin embedded tissues, some autofluorescence is expected which could explain some of reviewer’s concerns. However we expect that the inclusion of better quality images and quantification must address some of the concerns raised by the reviewer.

Eosinophils are heavily involved in lung macrophage biology, but are missing from the analysis - it is highly likely the RNA-sequence picked out eosinophils as Ly6G- neutrophils rather than 'digestion issues' the authors claim

We thank the reviewer for raising a valid concern. However, the Ly6G- cluster cannot be eosinophils in our case. Literature suggests SiglecF as an important biomarker of eosinophils which was absent in the Ly6G- cluster our in scRNA seq analyses as shown in File S18 and Figure 6B of the revised manuscript. We have now provided a detailed explanation (Lines 476-488; 503-506) of the observed results pertaining to eosinophil population in the revised manuscript to further address some of the concerns raised by this reviewer.

After author comments, it appears the schematic in Figure 1A is misleading and there are not n=2/group/sex but actually only n=1/group/sex (as shown in Figure 6A). Meaning the n number is even lower than the previous assumption.

We concur with reviewers’ valid concern and so are willing to provide this data as a resource for a wider audience to assist future work. Pooling of samples have been practiced by many groups previously to save resources and expense. We did it for the very same reason. It may not be the preferred approach, but it still has its merit considering the vast amount of cell-specific data generated using this strategy. To avoid overstating our results we have ensured to maintain transparency in our reporting and acknowledge all the limitations of this study.

We do not believe that the strength of scRNA seq lies in drawing conclusive results, but to tease our possible targets and direction that need to be validated with more work. In that respect, our study does identify the target cell types and biological processes which could be of importance for future studies.

Reviewer #3 (Public review):

This work aims to establish cell-type specific changes in gene expression upon exposure to different flavors of commercial e-cigarette aerosols compared to control or vehicle. Kaur et al. conclude that immune cells are most affected, with the greatest dysregulation found in myeloid cells exposed to tobacco-flavored e-cigs and lymphoid cells exposed to fruit-flavored e-cigs. The up- and down-regulated genes are heavily associated with innate immune response. The authors suggest that a Ly6G-deficient subset of neutrophils is found to be increased in abundance for the treatment groups, while gene expression remains consistent, which could indicate impaired function. Increased expression of CD4+ and CD8+ T cells along with their associated markers for proliferation and cytotoxicity is thought to be a result of activation following this decline in neutrophil-mediated immune response.

Strengths:

Single cell sequencing data can be very valuable in identifying potential health risks and clinical pathologies of lung conditions associated with e-cigarettes considering they are still relatively new.

Not many studies have been performed on cell-type specific differential gene expression following exposure to e-cig aerosols.

The assays performed address several factors of e-cig exposure such as metal concentration in the liquid and condensate, coil composition, cotinine/nicotine levels in serum and the product itself, cell types affected, which genes are up- or down-regulated and what pathways they control.

Considerations were made to ensure clinical relevance such as selecting mice whose ages corresponded with human adolescents so that data collected was relevant.

Weaknesses:

The exposure period of 1 hour a day for 5 days is not representative of chronic use and this time point may be too short to see a full response in all cell types. The experimental design is not well-supported based on the literature available for similar mouse models. Clinical relevance of this short exposure remains unclear.

We thank the reviewer for this query. However, we would like to emphasize that chronic exposure was never the intention of this study. We wished to design a study for acute nose-only exposure owing to which the study duration was left shorter. Shorter durations limit the stress and discomfort to the animal. The in vivo study using nose-only exposure regimen is still developing with multiple exposure regimen being used by different groups. To our knowledge there is no gold standard of e-cig aerosol exposure which is widely accepted other than the CORESTA recommendations, which we followed. Also, we show in our study how the daily exposure to leached metals vary in a flavor-dependent manner thus validating that exposure regime does need more attention in terms of equal dosing, particle distribution and composition- something we have started doing in our future studies. We have included all the explanations in the revised manuscript (Lines 82-85, 425-435, 648-654).

Several claims lack supporting evidence or use data that is not statistically significant. In particular, there were no statistical analyses to compare results across sex, so conclusions stating there is a sex bias for things like Ly6G+ neutrophil percentage by condition are observational.

We agree with reviewer’s comment and have taken this into consideration. We have now revamped the whole manuscript and toned down most of the sex-based conclusions stated in this work. Having said that, it is important to note that most of the work relying solely on scRNA seq, as is the case for this study, is observational in nature and needs to be assessed bearing this in mind.

Overall, the paper and its discussion are relatively surface-level and do not delve into the significance of the findings or how they fit into the bigger picture of the field. It is not clear whether this paper is intended to be used as a resource for other researchers or as an original research article.

We have now reworked on the Discussion and tried to incorporate more in-depth discussion and the results providing our insights regarding the observations, discrepancies and the possible explanations. We have also made it clear that this paper is intended to be used as a resource by other researchers (Lines 577-579)

The manuscript has some validation of findings but not very comprehensive.

We have now revamped the manuscript. We have Included quantification for immunofluorescence data with better representation of the GO analyses. We have worked on the Results and Discussion sections to make this a useful resource for the scientific community.

This paper provides a strong foundation for follow-up experiments that take a closer look at the effects of e-cig exposure on innate immunity. There is still room to elaborate on the differential gene expression within and between various cell types.

We thank the reviewer for pointing out the strength of this paper. The reason why we refrained from elaborating of the differential gene expressions within and between various cell types was due to low sample number and sequencing depth for this study. However the raw data will be provided with the final publication, which should be freely accessible to the public to re-analyze the data set as they deem fit.

Comments on revisions:

The reviewers have addressed major concerns with better validation of data and improved organization of the paper. However, we still have some concerns and suggestions pertaining to the statistical analyses and justifications for experimental design.

We appreciate the nuance of this experimental design, and the reviewers have adequately commented on why they chose nose-only exposure over whole body exposure. However, the justification for the duration of the exposure, and the clinical relevance of a short exposure, have not been addressed in the revised manuscript.

We thank the editor for this query. We have now addressed this query briefly in Lines 82-85, 425-435, 648-654 of the revised manuscript. We would like to add, however, that we intend to design a study for acute nose-only exposure for this project. Shorter durations limit the stress and discomfort to the animal, owing to which a duration of 1hour per day was chosen. The in vivo study using nose-only exposure regimen is still developing with multiple exposure regimen being used by different groups. Ours is one such study in that direction just intended to identify cell-specific changes upon exposure. Considering our results in Figure 1B showing variations in the level of metals leached in each flavor per day, the appropriate exposure regimen to design a controlled, reproducible experiment needs to be discussed. There could be room for improvement in our strategy, but this was the best regimen that we found to be appropriate per the literature and our prior knowledge in the field.

The presentation of cell counts should be represented by a percentage/proportion rather than a raw number of cells. Without normalization to the total number of cells, comparisons cannot be made across groups/conditions. This comment applies to several figures.

We thank the editor for this comment and have now made the requested change in the revised manuscript.

We appreciate that the authors have taken the reviewers' advice to validate their findings. However, we have concerns regarding the immunofluorescent staining shown in Figure 4. If the red channel is showing a pan-neutrophil marker (S100A8) and the green channel is showing only a subset of neutrophils (LY6G+), then the green channel should have far less signal than the red channel. This expected pattern is not what is shown in the figure, with the Ly6G marker apparently showing more expression than S100A8. Additionally, the FACS data states that only 4-5% of cells are neutrophils, but the red channel co-localizes with far more than 4-5% of the DAPI stain, meaning this population is overrepresented, potentially due to background fluorescence (noise). In addition, some of the shapes in the staining pattern do not look like true neutrophils, although it is difficult to tell because there remains a lot of background staining. The authors need to verify that their S100A8 and Ly6G antibodies work and are specific to the populations they intend to target. It is possible that only the brightest spots are truly S100A8+ or Ly6G+.

We thank the editor for this comment and acknowledge that we may have made broad generalizations in our interpretation of our data previously. We have now revisited the data and quantified the two fluorescence for better interpretation of our results. We have also reassessed our conclusions from this data and reworded the manuscript accordingly. Briefly we believe that Ly6G deficiency could be an indication of the presence of immature neutrophils in the lungs. This is a common process of neutrophil maturation. An active neutrophil population has Ly6G and should also express S100A8 indicating a normal neutrophilic response against stressors. However, our results, despite some autofluorescence which is common with lung tissues, shows a marked decline in the S100A8+ cells in the lung of tobacco-flavored e-cig aerosol exposed mice as compared to air controls. We also do not see prominent co-localization of the two markers in exposed group thus proving a shift in neutrophil dynamics which requires further investigation. We would also like to mention here that S100A8 is predominantly expressed in neutrophils, but is also expressed by monocytes and macrophages, so that could explain the over-representation of these cells in our immunofluorescence results. We have now included this in the Discussion section (Lines 489- 538) of the revised manuscript.

Paraffin sections do not always yield the best immunostaining results and the images themselves are low magnification and low resolution.

We agree with the editor that paraffin sections may not yield best results, we have worked on the final figure to improve the quality of the displayed results and zoomed-in some parts of the merged image to show the differences in the co-localization patterns for the two markers in our treated and control groups for easier interpretation.

Please change the scale bars to white so they are more visible in each channel.

The merged image in Figure 6C now has a white scale bar.

We appreciate that this is a preliminary test used as a resource for the community, but there is interesting biology regarding immune cells that warrants DEG analysis by the authors. This computational analysis can be easily added with no additional experiments required.

We thank the editor for this comment and agree that interesting biology regarding immune cells could be explored upon performing the DEG analyses on individual immune populations. However, due to the small sample size, low sequencing depth and pooling of same sex animals in each treatment group, we refrained from performing that analyses fearing over-representation of our results. We will be providing the link to the raw data with this publication which will be freely accessible to public on NIH GEO resource to allow further analyses on this dataset by the judgement of the investigator who utilizes it as a resource.

Recommendations for the authors:

Reviewer #1 (Recommendations for the authors):

(Minor) The pathway analyses in Fig. 6-8 have different fonts than what's used in all other figures.

We have now made the requested change in the revised manuscript.

Changes from the earlier version to the revised version

Abstract

Removed in the revised version: Abstract

Earlier version: E-cigarettes (e-cigs) are a public health concern for young adults due to their rising popularity despite evidence of harmful effects, i.e. increased oxidative stress and immunotoxicity. Yet an extensive study defining the cell-specific immune changes upon exposure to flavored e-cigs remains elusive. To determine the immunological lung landscape upon acute nose-only exposure of C57BL/6J to flavored e-cig aerosols, we performed single-cell RNA sequencing (scRNA seq). scRNA profiles of 71,725 cells were generated from control and treatment groups (n=2/sex/group). An increase in the Ly6G- neutrophil population was identified in lungs exposed to tobacco flavored e-cig aerosol. Differential gene expression analyses identified dysregulation of T-cell mediated pro-inflammation (Cct7, Cct8) and stress-response signals (Neurl3, Stap1, Cirbp and Htr2c) in the lungs of mice exposed to e-cig aerosols, with pronounced effects for tobacco flavor. Flow cytometry analyses and cytokine/chemokine assessments within the lungs corroborate the scRNA seq data, demonstrating an increase in the T-cell percentages and levels of T-cell associated cytokine/chemokines in the lungs of tobacco-flavored aerosol exposed mice. Increased levels of Klra4 and Klra8 expression also suggest an enhanced natural killer (NK) cell activity in this mouse group. Interestingly, the increase in the level of Ly6G – (immature) neutrophils upon exposure to tobacco-flavored e-cig aerosols was noticed for female C57BlL6J mice and not their male counterparts which was validated through immunofluorescence staining. Overall, this study identifies sex-specific increase in the percentages of Ly6G- neutrophils that may be responsible for dampened innate immune responses in females and heightened T-cell cytotoxicity in male mouse lungs of tobacco-flavored e-cig aerosol exposed mice.

Revised version: E-cigarettes (e-cigs) are a public health concern for young adults due to their rising popularity despite evidence of harmful effects. Yet an extensive study defining the cell-specific immune changes upon exposure to flavored e-cigs remains elusive. To determine the immunological lung landscape upon acute nose-only exposure of C57BL/6J to flavored e-cig aerosols, we performed single-cell RNA sequencing (scRNA seq). Analyses of the levels of metals in the e-cig aerosol generated daily during exposure revealed a flavor-dependent variation in the day-to-day leaching of the levels of metals like Ni, Cu, K and Zn, among others. scRNA profiles of 71,725 cells generated from control and treatment groups (n=2/sex/group) found maximum dysregulation of (a) myeloid cell function in menthol (324 differentially enriched genes (DEG)) and tobacco (553 DEGs) –flavor exposed and (b) lymphoid cell function in fruit (112 DEGs)-flavored e-cig aerosol exposed mouse lungs as compared to air. Flow cytometry analyses identified marked increase in the neutrophil percentage and a decrease in the eosinophil count in menthol and tobacco-flavored e-cig aerosol exposed mouse lungs which corroborated with our scRNA seq data. We further found an: (a) increase in CD8+ T cell percentages, (b) upregulation of inflammatory genes like Stat4, Il1bos, Il1b, Il1ra and Cxcl3 and (c) enrichment of terms like ‘T-helper cell 1cytokine function’ and ‘NK cell degranulation’ in the lungs of e-cig aerosol exposed mouse when compared to control. Interestingly, the increase in the level of immature neutrophils characterized by Ly6G deficiency and reduction in the S100A8 (marker for neutrophil activation) using immunofluorescence in tobacco-flavored e-cig aerosol exposed mouse lung sections point towards a possible shift in the neutrophil dynamics upon exposure to e-cigs. Overall, this study identifies flavor dependent changes in myeloid and lymphoid-cell mediated responses in the mouse lungs exposed to acute nose-only e-cig aerosols and provides a resource to influence future research in select /specific cell types to understand the immunological implications of long-term use of e-cigs.

Introduction

Electronic cigarettes (e-cigs) or electronic nicotine delivery systems are a relatively novel set of tobacco/nicotine and flavored products that have gained immense popularity amongst adolescents and young adults in many countries including the United States (US), the United Kingdom (UK) and China. Flavors are one of the key features that make these products alluring to the younger diaspora (1). Reports indicate that in 2020 about 22.5% of high school students and 9.4% of middle school students in the US were daily vapers or e-cig users with fruit (66%), mint (57.5%), and menthol (44.5%) being the most commonly used flavors (2). However, not much is known about the flavor-specific effects of e-cig vaping on the health and immunity of an individual specially focusing on cell-types and gene transcripts.

E-cig products and aerosols are known to contain harmful constituents including formaldehyde, benzaldehyde, acrolein, n-nitrosamines, volatile organic compounds (VOCs), ketenes, and metal ions (3–5). Studies have indicated that exposure to e-cigs may enhance inflammatory responses, oxidative stress, and genomic instability in exposed cells or animal systems (6–8). Risk assessment (systemic) of inhaled diacetyl, a potential component of e-liquids, has estimated the non-carcinogenic hazard quotient to be greater than 1 amongst teens (9). Furthermore, clinical and in vivo studies have suggested that exposure to e-cig aerosols could impair innate immune responses in the host thus making them more susceptible to bacterial/viral infections. The bacterial clearance, mucous production, and phagocytic responses in these individuals are shown to be affected upon use of e-cigs (10–14).

Earlier version: However, cell-specific changes within the lung upon vaping is not fully understood, making it hard to determine the long-term health impacts of the use of these novel products. In this respect, single cell technology is a powerful tool to analyze gene expression changes within cell populations to study cell heterogeneity and function (15–17). Such an investigation is important to deduce the health effects of acute and chronic use if e-cigs in young adults. In this study, we aim to determine effects of e-cig exposure on mouse lungs at single cell level. To do so, we exposed C57BL/6J mice to 5-day nose- only exposure to air, propylene glycol: vegetable glycerin (PG:VG), fruit-, menthol- and tobacco-flavored e-cig aerosols. The nose-only exposure has more translational relevance over whole-body exposure (18), owing to which, we chose nose-only exposure profile for this work. To limit the stress to the animals a 1-hour exposure was chosen per day. We performed single-cell RNA sequencing (scRNA seq) on the lung digests from exposed and control animals and identified neutrophils and macrophages, among others, as the major cell populations in the lung that were affected upon acute exposure. We were able to identify 29 gene targets that were commonly dysregulated amongst all our treatment groups upon aggregating results from the major lung cell types identified in this study. These gene targets are the markers of early immune dysfunction upon e-cig aerosol exposure in vivo and could be studied in detail to understand temporal changes in their expression that may govern allergic responses and adverse pulmonary health outcomes upon acute and sub-acute exposures to e-cigarette aerosols.

Revised version: However, cell-specific changes within the lung upon vaping is not fully understood, making it hard to determine the long-term health impacts of the use of these novel products. In this respect, single cell technology is a powerful tool to analyze gene expression changes within cell populations to study cellular heterogeneity and function (15–17). Such an investigation is important to deduce the health effects of acute and chronic use of e-cigs in young adults. In this study, we aim to determine the effects of e-cig exposure on mouse lungs at single cell level. To do so, we exposed C57BL/6J mice to 5-day nose-only exposure to air, propylene glycol: vegetable glycerin (PG:VG), fruit-, menthol– and tobacco-flavored e-cig aerosols. The nose-only exposure has more translational relevance over the whole-body exposure (18), owing to which, we chose nose-only exposure profile for this work. To limit the stress to the animals a 1-hour exposure was chosen per day. We performed single-cell RNA sequencing (scRNA seq) on the lung digests from exposed and control animals and identified neutrophils and T-cells, among others, as the major cell populations in the lung that were affected upon acute exposure. We were able to identify 29 gene targets that were commonly dysregulated amongst all our treatment groups upon aggregating results from the major lung cell types. These gene targets are the markers of early immune dysfunction upon e-cig aerosol exposure in vivo and could be studied in detail to understand temporal changes in their expression and function that may govern allergic responses and adverse pulmonary health outcomes upon acute and sub-acute exposures to e-cigarette aerosols.

Results

Earlier version: Exposure to e-cig aerosol results in flavor-dependent exposure to different metals and mild histological changes in vivo

Revised version: Exposure to e-cig aerosols results in flavor-dependent exposure to different metals and mild histological changes in vivo

Earlier version: This study was designed to characterize the effects of exposure to flavored e-cig aerosols at single-cell level to understand the immunological changes in the lung microenvironment upon e-cig use. To do so, we generated a first ever single-cell profile of e-cig aerosol exposed mouse lungs (n= 2/sex/group). The thus obtained results were then validated with the help of our validation cohort of n= 3/sex/group as shown in Fig 1A. Since all the commercially available e-liquids used in this study contained tobacco derived nicotine, we first determined the level of serum cotinine (a metabolite of nicotine) to prove successful exposure to the mice in each treatment group. As expected, we did not see any traces of cotinine in the serum of air and PG:VG exposed mice. Significant levels of cotinine were detected in the serum of mice exposed to fruit-, menthol-, and tobacco-flavored e-cig aerosols (Figure S1B).

Revised version: This study was designed to characterize the effects of exposure to flavored e-cig aerosols at single-cell level to understand the immunological changes in the lung microenvironment. To do so, we generated a single-cell profile of e-cig aerosol exposed mouse lungs (n= 2/sex/group). The thus obtained results were then validated with the help of our validation cohort of n= 3/sex/group as shown in Figure 1A. Since all the commercially available e-liquids used in this study contained tobacco derived nicotine (TDN), we first determined the levels of serum cotinine (a metabolite of nicotine) to prove successful exposure of the mice in each treatment group. As expected, we did not see any traces of cotinine in the serum of air and PG:VG exposed mice. Significant levels of cotinine were detected in the serum of mice exposed to fruit-, menthol-, and tobacco-flavored e-cig aerosols (Figure S1B), validating successful exposure of our test animals to TDN containing flavored e-cigs.

Earlier version: Flavor-dependent changes in the levels of quantified metals, but no major histological damage on acute exposure to flavored e-cig aerosol in mice.

Revised version: Flavor-dependent changes in the levels of quantified metals, but no major histological damage on acute exposure to flavored e-cig aerosol in C57BL/6J mice.

Earlier version: Schematics showing the exposure profile and experimental design to understand the effect of exposure to differently flavored e-cig aerosols in the lungs from C57BL/6J mice using scRNA seq (A)). Bar Graph showing the levels of metals in the aerosols captured each day of exposure using ICP-MS (B). Lung morphometric changes as observed using H&E staining of lung slices from air, PG:VG and differently flavored e-cig exposed mice. Representative images of n = 2/sex/group at 10X magnification is provided (C).

Revised version: Schematics showing the exposure profile and experimental design to understand the effects of exposure to differently flavored (fruit, menthol and tobacco) e-cig aerosols in the lungs of C57BL/6J mice using scRNA seq (A). Bar Graph showing the levels of metals as determined by ICP-MS in the aerosols captured daily during exposures using inExpose nose-only inhalation system from Scireq technologies (B). Lung morphometric changes observed using H&E staining of lung slices from air, PG:VG and differently flavored e-cig aerosol exposed mice lungs. Representative images of n = 2/sex/group at 10X magnification is provided (C).

Earlier version: Since metals released upon heating of the coils of e-cig devices are a source of toxicity upon vaping (19, 20), we further monitored the levels of metals in the e-cig aerosols generated during each day of mouse exposures. This acted as an indirect measure of characterizing the chemical properties of the aerosols used for exposure in this study. To monitor the release of metals into the lungs of the animals, the aerosol condensate from each day of exposure was collected and the levels of select elements were detected using ICP-MS. A detailed account of the concentrations of identified elements/metals is provided in Table 1. Interestingly, we identified flavor-dependent changes in the levels of metals like Ni, Zn, Na, K and Cu. Note, despite the use of same wattage and temperature (max of 230⁰C) for generation of e-cig aerosols, the leaching of each metal varied per day of exposure (Figure 1B). This is a crucial result as it highlights the importance of studying the impacts of coil composition and device design on the chemical composition of generated e-cig aerosols. These variations might affect the risk and toxicity associated with each of these products, an area that warrants further study.

Revised version: Since metals released upon heating of the coils of e-cig devices are a source of toxicity upon vaping (19, 20), we further monitored the levels of metals in the e-cig aerosols generated during each day of mouse exposures. This acted as an indirect measure for characterizing the chemical properties of the aerosols used for exposure in this study. To monitor the release of metals into the mouse lungs, the aerosol condensate from each day of exposure was collected and the levels of select elements were detected using ICP-MS. A detailed account of the concentrations of identified elements/metals is provided in Table 1. Interestingly, we identified flavor-dependent changes in the levels of metals like Ni, Zn, Na, K and Cu on a day-to-day basis. Note, despite the use of same wattage and temperature (max of 230 C) for generation of e-cig aerosols, the leaching of each metal varied per day of exposure (Figure 1B). This is a crucial result as it highlights the importance of studying the impact of atomizer, coil composition and design on the chemical composition of the generated aerosols. These variations might affect the risk and toxicity associated with each of these products, an area that has been recently explored by our group (21).

Earlier version: Next, we performed H&E staining of the lung tissue sections to study the morphometric changes in the mouse lungs upon exposure to differently flavored e-cig aerosols. We did not find much evidence of tissue damage or airspace enlargement upon acute exposures in our model, as expected. However, we found evidence of increased alveolar septa thickening in the lungs of both male and female mice exposed to fruit-flavored e-cig aerosols (Figure 1C). However, this could be a result of lower agarose inflation observed for the mouse lungs in this group. We had challenges/difficulties with inflating the mouse lungs for all mice in this group, however male mice showed more restriction than the female. Owing to the lack of proper inflation, the lungs were in a collapsed state than the other groups which could be a probable cause of the histological observations. Airway obstruction could be indicative of “atelectasis”, a condition that may result from lack of surfactant or alveolar damage thus affecting the lung expansion in these mice. Since this was not the prime focus of this study, we did not conduct further experiments to confirm our speculations, but nevertheless, we report these changes to inform and encourage further research in this area.

Revised version: Next, we performed H&E staining on the lung tissue sections to study the morphometric changes in the mouse lungs upon exposure to differently flavored e-cig aerosols. We did not find much evidence of tissue damage or airspace enlargement upon acute exposures in our model, as expected. However, we found evidence of increased alveolar septa thickening in the lungs of both male and female mice exposed to fruit-flavored e-cig aerosols (Figure 1C). This could be a result of lower agarose inflation observed for the mouse lungs in this group. We had challenges/difficulties with inflating these mouse lungs. Owing to the lack of proper inflation, the lungs were in a collapsed state which could be a probable explanation for our histological observation. Since this was not the prime focus of our study, we did not conduct further experiments to confirm our speculations.

Detailed map of cellular composition during acute exposure to e-cig aerosols reveals distinct changes in immune cell phenotypes

Earlier version: As the principle focus of this study was to identify the flavor-dependent and independent toxicities upon exposure to commercially available e-cig aerosols at the single-cell level, we performed the scRNA seq on the mouse lungs from exposed and control mice. After quality control filtering (Figure S2), normalization and scaling, we generated scRNA seq profiles of 71,725 cells in total. Except for the PG:VG group, all the rest of the treatments had approximately similar cell viabilities, cell capture, and other quality assessments in this study. However, for uniform analyses, equal features/ genes were used across all the groups for subsequent analyses. A detailed account of the cell viability, number of samples (single-cell capture), and gene features identified before and after filtering upon QC check of sequenced data is provided in Supplementary File S1.

Revised version: As the principle focus of this study was to identify the flavor-dependent and independent effects upon exposure to commercially available e-cig aerosols at the single-cell level, we performed the scRNA seq on the mouse lungs from exposed and control mice. After quality control filtering (Figure S2), normalization and scaling, we generated scRNA seq profiles of 71,725 cells in total. Except for the PG:VG group, all the rest of the treatments had approximately similar cell viabilities, cell capture, and other quality assessments. However, for normalization, equal features/ genes were used across all the groups for subsequent analyses. A detailed account of the cell number (single-cell capture) and gene features identified before and after filtering upon QC check of sequenced data is provided in Supplementary File S1.

Earlier version: Uniform Manifold Approximation and Projection (UMAP) was used for dimensionality reduction and visualization of cell clusters. Cell annotations were performed based on the established cell markers in Tabula Muris database and available published literature and we identified 24 distinct cell clusters as shown in Figure 2A. The general clustering of individual cell types based upon the commonly known cell markers was used to identify -Endothelial (identified by expression of Cldn5), Epithelial (identified by expression of Sftpa1), Stromal (identified by expression of Col3a1), and Immune (identified by expression of Ptprc) cells (Figure 2B). The ‘FindVariableFeatures’ from Seurat was used to identify cell-to-cell variation between the identified clusters and the top 2000 variable genes identified in each cluster have been elaborated in Supplementary File S2. Upon plotting the number of counts/cells identified in each cluster in a sex- and treatment-specific manner we observed minor changes, though not significant, in the count of myeloid and lymphoid cell clusters in e-cig aerosol exposed mouse lungs as compared to air controls (Figure 2C). A detailed account of the two-way ANOVA statistics for individual cell types is provided in Supplementary File S3A.

Revised version: Uniform Manifold Approximation and Projection (UMAP) was used for dimensionality reduction and visualization of cell clusters. Cell annotations were performed based on the established cell markers in Tabula Muris database and available published literature and we identified 24 distinct cell clusters as shown in Figure 2A. The general clustering of individual cell types based upon the commonly known cell markers was used to identify –Endothelial (identified by expression of Cldn5), Epithelial (identified by expression of Sftpa1), Stromal (identified by expression of Col3a1), and Immune (identified by expression of Ptprc) cell populations (Figure 2B). The ‘FindVariableFeatures’ from Seurat was used to identify cell-to-cell variation between the identified clusters and the top 2000 variable genes identified in each cluster have been elaborated in Supplementary File S2. We observed minor variations in the cell frequencies as observed through scRNA seq analyses within cell types across different treatment groups (Figure S3A). The largest proportion of cells was found in the endothelial cell cluster (43.87% of total per sample), followed by lymphoid (26.43% of total per sample) and myeloid (16.79% of total per sample) clusters. A detailed account of the two-way ANOVA statistics for the general clustering with cell types and treatment groups as independent variables is provided in Supplementary File S3A.

Earlier version: scRNA seq analyses reveal maximum changes in the cell states of immune cell population upon exposure to differently flavored e-cig aerosols.

Revised version: scRNA seq analyses reveal maximum changes in the transcriptional profile of immune cell population upon exposure to differently flavored e-cig aerosols.

Earlier version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to differently flavored e-cig aerosols. The mice were sacrificed after the final exposure and lungs from air (control) and differently flavored e-cig aerosol (fruit, menthol, and tobacco)-exposed mice were used to perform scRNA seq. UMAP plot of 71,725 cells captured during scRNA seq showing the four major cell populations identified from control and experimental mouse lungs (A) and the expression of canonical markers used for identifying stromal (Col3a1), epithelial (Sftpa1), endothelial (Cldn5) and immune (Ptprc) cell populations. The intensity of expression is indicated by the red-yellow coloring (B). The cell frequencies (plotted as counts) of different cell clusters in each sample type showing the sex-dependent variations in the cell composition on exposure to differently flavored e-cig aerosols (C). Group-wise comparison of the UMAPs upon comparing PG:VG (blue), fruit (yellow), menthol (green) and tobacco (red) versus air (air) groups following dimensionality reduction and clustering of scRNA seq data (D). Here, AT1: alveolar type I, AT2: alveolar type II, SMC: smooth muscle cell, gCap: general capillary, aCap: alveolar capillary, and NK: natural killer.

Revised version: Male and female C57BL/6J mice (n = 2/sex/group) were exposed to 5-day nose-only exposure to differently flavored e-cig aerosols. The mice were sacrificed after the final exposure and lungs from air (control) and differently flavored e-cig aerosol (fruit, menthol, and tobacco)-exposed mice were used to perform scRNA seq. UMAP plot of 71,725 cells captured during scRNA seq showing the 24 major cell clusters identified from control and experimental mouse lungs (A) and the expression of canonical markers used for identifying stromal (Col3a1), epithelial (Sftpa1), endothelial (Cldn5) and immune (Ptprc) cell populations. The intensity of expression is indicated by the black-yellow coloring (B). Group-wise comparison of the UMAPs upon comparing PG:VG (blue), fruit (yellow), menthol (green) and tobacco (red) versus air (grey) groups following dimensionality reduction and clustering of scRNA seq data (C). Bar plot showing the number of significant (p < 0.05) differentially up-(green) and downregulated (red) genes in myeloid and lymphoid clusters in PGVG, fruit-, menthol-and tobacco-flavored e-cig aerosol exposed mouse lungs as compared to air controls. Here, AT1: alveolar type I, AT2: alveolar type II, Fibro: Fibroblast, MΦ: macrophage, SMC: smooth muscle cell, gCap: general capillary, aCap: alveolar capillary, and NK: natural killer.

Earlier version: Groupwise comparisons of each treatment group versus Air control, did not show many changes in the cell clusters thus showcasing little to no effect on the overall lung cellular compositions in treated and control groups. However, the number of cells identified in each cluster for PG:VG treatment group was very low. We would like to report that this is an outcome of the low viability of this sample prior to scRNA seq, and not necessarily the effect of the treatment (Figure 2D). Differential gene expression analyses showed dysregulation of genes in all cell populations, but maximum effect was observed on the cells from immune cell clusters. This is not surprising, as the immune system, especially the myeloid cells form the frontline of the host’s defenses against external stressors (21, 22). Compared to air, we observed a dysregulation of 553 genes (338 upregulated; 215 downregulated) in the myeloid cell cluster of mouse lungs exposed to tobacco-flavored e-cig aerosol. We identified 324 and 24 DEGs in the myeloid lung cluster from mouse lungs exposed to menthol- and fruit-flavored e-cig aerosols respectively as compared to air control (Supplementary Files S5-S7). For the lymphoid cluster, we observed maximum dysregulation in the lungs exposed to fruit-flavored e-cig aerosols with a total of 112 DEGs. In contrast, 41 and 9 significant DEGs were identified for lymphoid cluster from lungs exposed to tobacco and menthol-flavored aerosols respectively (Supplementary Files S9-S11).

Revised version: Groupwise comparisons of each treatment group versus air control, did not show many changes in the cell clusters thus showcasing little to no effect on the overall lung cellular compositions in treated and control groups. However, the number of cells for PG:VG group was very low (8102 cells) as compared to other treatments (∼15,710 cells on average). We would like to report that this is an outcome of the low viability of this sample prior to scRNA seq, and not necessarily the effect of the treatment (Figure 2C). Differential gene expression analyses showed dysregulation of genes in all cell populations, but maximum effect was observed on the cells from immune cell clusters. This is not surprising, as the immune system, especially the myeloid cells form the frontline of the host’s defenses against external stressors (22, 23). Compared to air, we observed a dysregulation of 553 genes (338 upregulated; 215 downregulated) in the myeloid cell cluster of mouse lungs exposed to tobacco-flavored e-cig aerosol. We identified 324 and 24 DEGs in the myeloid lung cluster from mouse lungs exposed to menthol– and fruit-flavored e-cig aerosols respectively as compared to air control (Figure 2D; Supplementary Files S5-S10). For the lymphoid cluster, we observed maximum dysregulation in the lungs exposed to fruit-flavored e-cig aerosols with a total of 112 DEGs. In contrast, 41 and 9 significant DEGs were identified for lymphoid cluster from lungs exposed to tobacco and menthol-flavored aerosols respectively (Figure 2D; Supplementary Files S12-S17).

Earlier version: It is important to mention here that most of the commercially available e-liquids /e-cig products use PG:VG as the base liquid to generate the aerosol and act as a carrier for flavoring chemicals. Thus, we compared the effect of PG:VG alone in our study to make further comparisons between individual flavoring products with that of PG:VG only. DESeq2 analyses showed dysregulation of 276 genes in mouse lungs exposed to PG:VG alone in the myeloid cell cluster as compared to air controls (Supplementary Files S4). Contrary to this, exposure to PG:VG aerosols affected 24 genes in the lymphoid cluster as compared to air control (Supplementary Files S8). Furthermore, when compared to PG:VG, exposure to fruit, menthol and tobacco-flavored e-cig aerosol dysregulated 262, 873 and 960 genes in the myeloid cluster and 37, 64 and 112 genes in the lymphoid cluster respectively. A detailed account of the DEGs identified upon comparing fruit, menthol and tobacco-flavored e-cig aerosol exposed mouse lungs to PG:VG in myeloid and lymphoid clusters have been provided in Supplementary Files S13-S18.

Revised version: It is important to mention here that most of the commercially available e-liquids /e-cig products use PG:VG as the base liquid to generate the aerosol and act as a carrier for flavoring chemicals. Thus, we compared the effect of PG:VG alone in our study to make further comparisons between individual flavoring products with that of PG:VG only. DESeq2 analyses showed dysregulation of 276 genes in mouse lungs exposed to PG:VG alone in the myeloid cell cluster as compared to air controls (Figure 2D; Supplementary Files S4). Contrary to this, exposure to PG:VG aerosols affected 24 genes in the lymphoid cluster as compared to air control (Figure 2D; Supplementary Files S11). Furthermore, when compared to PG:VG, exposure to fruit, menthol and tobacco-flavored e-cig aerosol dysregulated 262, 873 and 960 genes in the myeloid cluster and 37, 64 and 112 genes in the lymphoid cluster respectively. A detailed account of the DEGs identified upon comparing fruit, menthol and tobacco-flavored e-cig aerosol exposed mouse lungs to PG:VG in myeloid and lymphoid clusters have been provided in Supplementary Files S19-S24.

Earlier version: Exposure to differently flavored e-cig aerosols results in sex-specific variations in neutrophilic and eosinophilic response in the mouse lungs

Revised version: Exposure to e-cig aerosols results in flavor-dependent changes in neutrophilic and eosinophilic response in vivo

Added in the revised version: To study the specific changes in the innate and adaptive immunity upon acute exposure to differently flavored e-cig aerosols, we first compared the changes in the overall cell population of individual cell types using scRNA seq. The scRNA seq data was validated with the help of flow cytometry using a larger cohort of animals. Since we observed sex-dependent variations in oxidative stress responses upon exposure to e-cig aerosols in previous studies (24–26), the flow cytometry data was analyzed in a sex-dependent fashion.

Added in the revised version: We did not observe any changes in the cell frequencies of alveolar macrophages across treatments through scRNA seq analyses. In general, there was a moderate decrease in the cell frequencies of alveolar macrophages in exposed mice as compared to air control independent of the flavor profile of the e-liquid employed (Figure 3A). Flow cytometric analyses confirmed little to no change in the alveolar macrophage (CD45+ CD11b− SiglecF+) percentages within the lung of exposed versus control mice (Figure 3C-D; Supplementary File S3B). Contrarily, we observed a flavor-dependent moderate increase in the cell frequencies of neutrophil cluster in menthol (0.1 ± 0.02) %– and tobacco (0.1 ± 0.06) %-flavored aerosol exposed mouse lung as compared to controls (0.06 ± 0.02) % (Figure 3B). We used flow cytometry to validate the scRNA seq changes. Flow cytometry analyses showed an increase in the neutrophil (CD45+ CD11b+ Ly6G+) percentages of menthol-flavored e-cig aerosol exposed mouse lung corroborating with the scRNA seq results. Furthermore, our results showed this increase to be more pronounced in male mice (p = 0.0880) as compared to their female (p = 0.9662) counterparts (Figure 3C-D, Supplementary File S3B).

Removed in the revised version: We found dysregulation in immune cell population, more specifically the myeloid cells, upon exposure to e-cig aerosols through our previous analyses. We thus plotted the changes in cell counts across the various treatment groups and control in a sex-dependent manner. scRNA seq results did not show any change in the cell frequencies of alveolar macrophages in flavored e-cig aerosol exposed mouse lungs as compared to controls (Figure 3A; Supplementary File S3B). However, the cell count of neutrophils in tobacco-flavored e-cig aerosol exposed mice lungs were found to be increased in females, but not in males (Figure 3B, Supplementary File S3B). Taken together, we show a significant shift towards the neutrophil-mediated immune response in the lungs of female mice exposed to tobacco-flavored e-cig aerosols. We used flow cytometry to validate the scRNA seq changes. While the changes observed in the levels of alveolar macrophages in the mouse lungs corroborated with our scRNA seq findings, the results for the neutrophil cell cluster did not show matching trend. This could be due to the increase in the levels of immature neutrophils in the lungs of female tobacco-e-cig exposed mouse lungs characterized by the absence of Ly6G cell surface marker, which was used to stain for neutrophils via flow cytometry (Figure 3C-D, Supplementary File S3D).

Earlier version: Cellular composition of myeloid cells in air and e-cig aerosol exposed mouse lung reveal sex-specific increase in neutrophil count through scRNA seq but not flow cytometry.

Revised version: Cellular composition of myeloid cells in air and e-cig aerosol exposed mouse lungs reveal increase in neutrophil count through scRNA seq and flow cytometry.

Earlier version: Cell frequencies (represented as counts) of alveolar macrophages (A) and neutrophils (B) across controls and flavored e-cig aerosol exposed mouse lungs were plotted in a sex-specific manner. Representative flow plots (C) and bar graphs (D) showing sex-dependent changes in the percentages of neutrophils (CD45+ CD11b+ Ly6G+) and alveolar macrophages (CD45+ CD11b-SiglecF+) populations in lung digests from mice exposed to differently flavored e-cig aerosols. Data are shown as mean ± SEM (n = 3/sex/group).

Revised version: Relative cell frequencies of alveolar macrophages (A) and neutrophils (B) across controls and flavored e-cig aerosol exposed mouse lungs as determined using scRNA seq. Representative flow plots (C) and bar graphs (D) showing sex-dependent changes in the percentages of neutrophils (CD45+ CD11b+ Ly6G+) and alveolar macrophage (CD45+ CD11b-SiglecF+) populations in lung digests from mice exposed to differently flavored e-cig aerosols. Values plotted and written in red on the flow plots are representative of the percentage of each cell population in the total CD45+ cells present in the lung homogenates from treatment and control groups. Data are shown as mean ± SEM (n = 3/sex/group). SE determined using two-way ANOVA with a Tukey post-hoc test for all cell means, to analyze the main effects of sex and treatment and their interaction. The two-way ANOVA results are shown in Supplementary File S3B.

Added in the revised version: It is important to mention that due to the presence of nicotine in all the e-liquids used as treatments for this study, we added an extra control group-PG:VG+Nic-for select experiments. Flow cytometric analyses revealed slight variations in the lung neutrophils and macrophage percentages observed in PG:VG+Nic group as compared to PG:VG only. However, none of these changes were significant. Furthermore, the patterns of change observed for the lungs exposed to aerosols from flavored e-liquids were quite distinct from those observed for PG:VG+Nic thus proving that the observed changes are not solely due to the presence of nicotine in these groups (Figures S3B).

Added in the revised version: Though we did not find a distinct eosinophil cluster in our scRNA seq analyzed data, interesting sex– and flavor-specific changes were observed in the lung eosinophil (CD45+ CD11b+/- CD11c− Ly6G− SiglecF+) population upon analyzing the flow cytometry results. A significant decline in the eosinophil percentage was observed in the lung homogenates from menthol (p =0.0113) and tobacco (p = 0.0312)-flavored e-cig aerosol exposed male C57BL/6J mice as compared to air control. Contrarily, the eosinophil levels in the lung of male mice exposed to fruit flavored e-cig aerosols did not show remarkable changes when compared to the levels found in control (Figure 4 & S4A). These results further emphasize the need of studying the sex-specific changes in the immune responses upon exposure to e-cig aerosols in vivo.

Added in the revised version: Flow cytometry analyses show significant decrease in the percentage of eosinophils in the lungs of menthol and tobacco-flavored e-cig aerosol exposed C57BL/6J mice.

Added in the revised version: Representative flow plot (A) and bar graphs (B) showing the changes in the percentages of eosinophils (CD45+ CD11b+/- CD11c− Ly6G− SiglecF+) found in the lungs of differently flavored e-cig aerosol exposed mouse lungs as compared to air controls. Data are shown as mean ± SEM (n = 3/sex/group). *p<0.05, per two-way ANOVA with a Tukey post-hoc test for all cell means, to analyze the main effects of sex and treatment and their interaction. The two-way ANOVA results are shown in Supplementary File S3B. Values plotted and written in red on the flow plots are representative of the percentage of each cell population in the total CD45+ cells present in the lung homogenates from treatment and control groups.

Added in the revised version: Activation of T-cell cytotoxic responses in lymphoid cells upon exposure to tobacco-flavored e-cig aerosols

Added in the revised version: We next studied the changes in the lymphoid clusters of treated and control samples. While we did not notice any change in the CD4+ T cell frequencies in treatment and control groups using scRNA seq, our flow cytometry data showed significant increase in the CD4+ T-cell (CD45+ CD11c− Ly6G− CD11b− MHCII− CD4+) frequencies in the lungs of tobacco-flavored e-cig aerosol exposed female (p= 0.0492) C57BL/6J mice (Figures 5A, 5C-D). scRNA data did not show any change in the cell frequencies of CD8+ T cells in control and exposed mouse lungs either (Figures 5B). However, flow cytometric analyses presented a different picture. We found a significant sex-independent increase in the CD8+ T (CD45+ CD11c− Ly6G− CD11b−MHCII−CD8+) cell percentages in the lungs of menthol– and tobacco-flavored e-cig aerosol exposed mice as compared to air control. Contrarily, the CD8+ T cell percentages increased in the fruit-flavored e-cig aerosol exposed male (p= 0.0163) mouse lungs and not in their female counterparts (Figure 5C-D, Supplementary File S3B). Of note, we did not observe any changes in the CD4+ and CD8+ T cell percentages in PGVG+Nic group as compared to control, reiterating that the e-cig flavors are responsible for altered immune responses upon acute exposure in mice (Figure S4B-C).

Added in the revised version: Flow cytometry results show sex-specific flavor-dependent increase in CD8+ T cells in lungs of differently flavored e-cig aerosol exposed C57BL/6J mouse.

Added in the revised version: Relative cell frequencies of CD4+ (A) and CD8+ (B) T-cells across controls and flavored e-cig aerosol exposed mouse lungs as determined using scRNA seq. Representative flow plots (C) and bar graph (D) showing changes in the mean cell percentages of CD4+ (CD45+ CD11c− Ly6G− CD11b− MHCII− CD4+) and CD8+ (CD45+ CD11c− Ly6G− CD11b− MHCII− CD8+) T-cells in the lung tissue digests from male and female mice exposed to differently flavored e-cig aerosols as determined using flow cytometry. Values plotted and written in red on the flow plots are representative of the percentage of each cell population in the total CD45+ cells present in the lung homogenates from treatment and control groups. Data are shown as mean ± SEM (n = 3/sex/group). *p<0.05, **p<0.01 and ***p<0.001; per two-way ANOVA with a Tukey post-hoc test for all cell means, to analyze the main effects of sex and treatment and their interaction. The two-way ANOVA results are shown in Supplementary File S3B.

Added in the revised version: Sub-clustering of myeloid cluster identifies a population of neutrophils devoid of Ly6G

Earlier version: It is important to mention that due to the presence of nicotine in all the e-liquids used as treatment groups for this study, we added an extra control group- PG:VG+Nic- for select experiments. Flow cytometric analyses revealed slight variations in the lung neutrophils and macrophage percentages observed in PG:VG+Nic group as compared to PG:VG only. However, none of these changes were significant. Furthermore, the patterns of change observed for the lungs exposed to aerosols from flavored e-liquids were quite distinct from those observed for PG:VG+Nic thus proving that the observed changes are not solely due to the presence of nicotine in these groups (Figures S3A-B). To probe further, we subclustered the myeloid cell populations. Upon sub clustering, we identified 14 unique clusters comprising all the major cell phenotypes including neutrophils, alveolar macrophages (AM), interstitial macrophages (IM), monocytes, dendritic cells (DC), and mast cells (Figure S3C). A detailed account of all the cell types identified with their respective marker genes is provided in Supplementary File 12. On deeper evaluation, we identified two unique phenotypes of neutrophils (identified by S100a8, Cxcl2, Sell and Lcn2) in the mouse lungs. These clusters were named as Ly6G+ Neutrophils and Ly6G- Neutrophils based on the presence or absence of Ly6G marker, respectively (Figure S3C). Ly6G is important for neutrophil migration, maturation and function within the lung (23). Interestingly we found an increase in the levels of both Ly6G+ and Ly6G- neutrophil in the tobacco-flavored e-cig aerosol exposed female mouse lungs as compared to air control. Though we did not observe similar changes in male mice exposed to tobacco-flavored e-cig aerosol (Figure 4A-B, Supplementary File S3C). The scRNA seq findings were further validated using co-immunofluorescence using S100A8 (red, pan-neutrophil marker) and Ly6G (green, marker of mature neutrophil). Co-immunofluorescence results show increase in levels of both Ly6G + (yellow stained) and Ly6G – (red stained) neutrophil in female mice exposed to tobacco flavored e-cig aerosol, but not males further corroborating our previous findings and suggesting that sex-specific changes in innate immune response in vivo (Figures 4C-D).

Revised version: To probe further, we subclustered the myeloid cell populations. Upon sub clustering, we identified 14 unique clusters comprising all the major cell phenotypes including neutrophils, alveolar macrophages (AM), interstitial macrophages (IM), monocytes, dendritic cells (DC), and mast cells (Figure 6A). A detailed account of all the cell types identified with their respective marker genes in myeloid subcluster is provided in Supplementary File S18. On deeper evaluation, we identified two unique phenotypes of neutrophils (identified by S100a8, S100a9, Il1b, Retnlg, Mmp9 and Lcn2) in the mouse lungs. These clusters were named as ‘Ly6G+ Neutrophils’ and ‘Ly6G− Neutrophils’ based on the presence or absence of Ly6G marker, respectively (Figure 6B). Ly6G is important for neutrophil migration, maturation and function within the lung (27). Since we did not anticipate identifying such a population of neutrophils following acute e-cig exposure, we did not gather flow cytometry evidence to validate this finding. We speculated that the Ly6G-population is a population of immature neutrophils. To confirm this possibility, we performed co-immunofluorescence using S100A8 (red; marker for neutrophil activation) and Ly6G (green) specific antibodies for control (air) and tobacco-flavored e-cig aerosol exposed mouse lungs. Co-immunofluorescence results showed a moderate increase (p =0.4429) in the level of Ly6G+ cells in the lungs of tobacco-flavored aerosol exposed mice as compared to air. However, the level of S100A8+ cells decreased markedly (p = 0.0571) in e-cig exposed group, thus showing that activation of neutrophils could be affected upon acute exposure to e-cigs (Figures 6C-D; Supplementary Figure S4D).

Co-immunofluorescence validates the increase of Ly6G– and Ly6G+ neutrophil population in the tobacco flavored e-cig exposed female C57BL/6J mice.

Earlier version: The myeloid cell clusters were subsetted to identify two populations of neutrophils with and without the presence of Ly6G cell marker representing mature and immature neutrophils respectively (Figure S3C). Bar graph showing the sex-specific changes in the cell frequencies (denoted as counts) of Ly6G+ (A) and Ly6G- (B) neutrophils in the lung of tobacco-flavored e-cig aerosol exposed mouse lung as compared to air control. scRNA seq findings were validated by staining the tissue sections from tobacco-flavored e-cig aerosols and control (air) with Ly6G (green) and S100A8 (red, pan-neutrophil marker). Representative images showing the co-immunostaining of Ly6G and S100A8 (shown as yellow puncta) in female (C) and male (D) control and tobacco-flavored e-cig aerosol exposed mice. Here Ly6G+ neutrophils are represented by double staining using yellow arrows, while Ly6G- neutrophils are denoted by red arrows.

Revised version: The myeloid cell clusters were subsetted to identify two populations of neutrophils with and without the presence of Ly6G cell marker representing mature and immature neutrophils respectively. UMAP showing the 14 distinct cell populations identified upon subsetting and re-clustering the myeloid clusters from scRNA seq dataset from control and e-cig exposed mouse lungs (A). Marker plot showing the differential expression of highly expressed genes in the Ly6G+ and Ly6G-neutrophil cluster. The intensity of expression is indicated by the yellow-blue coloring; black represents nil value for expression for that gene (B). scRNA seq findings for presence of mature (Ly6G+) and immature (Ly6G-) neutrophils were validated by staining the tissue sections from tobacco-flavored e-cig aerosols and control (air) with Ly6G (green) and S100A8 (red, neutrophil activation marker). Representative images showing the co-immunostaining of Ly6G and S100A8 (shown as yellow puncta) (C) with respective quantification of relative fluorescence for Ly6G and S100A8 (D) in control and tobacco-flavored e-cig aerosol exposed mice. Data are shown as mean ± SEM (n = 4/group). SE calculated per Mann-Whitney U test for pairwise comparisons. Here, Neu: neutrophil, AM: alveolar macrophage, RAM: resident AM, pRAM: Proliferating RAM, IM: interstitial macrophage, CM: classical monocyte, NCM: non-classical monocyte, and DC: dendritic cell.

Removed in the revised version: Though we did not find a distinct eosinophil cluster through our scRNA seq analyses, a significant decline in the level of eosinophils in the lungs of female mice exposed to menthol-flavored e-cig aerosol was observed through flow cytometric analyses. Though not significant, marked decline in the eosinophil count was observed upon acute exposure of male and female C57BL/6J mice to tobacco-flavored e-cig aerosols. Contrarily, the eosinophil count in the lung of male mice exposed to fruit flavored e-cig aerosols was observed to be significantly higher as compared to PG:VG exposed lungs (Figures S4).

Removed in the revised version: Activation of T-cell cytotoxic responses in lymphoid cells upon exposure to tobacco-flavored e-cig aerosols

Removed in the revised version: We next studied the changes in the lymphoid clusters of treated and control samples. Both scRNA seq and flow cytometric analyses showed no variation in the CD4+ T cell frequency in the lungs of different e-cig flavored exposed mouse lungs as compared to air controls. We also identified flavor-dependent variations in the cell frequencies of CD8 T cells (p= 0.0802) in the lymphoid cluster from mouse lungs (Figures 5A-B, Supplementary File S3B). Flow cytometric analyses corroborated with the scRNA seq findings and showed a significant increase in the CD8+ T cell percentages in the lungs of tobacco-flavored e-cig aerosol exposed mice lungs as compared to air control in male mice (Figure 5C-D, Supplementary File S3D). Of note, we found a significant increase in the CD4+ T cell percentages in menthol flavored e-cig aerosol exposed mouse lung as compared to PG:VG exposed control (Figure S5).

Removed in the revised version: scRNA seq analyses and flow cytometry result show flavor-dependent increase in CD8 T cells in lungs of differently flavored e-cig aerosol exposed C57BL/6J mouse.

Removed in the revised version: The cell frequencies (denoted as counts) of CD4 (A) and CD8 (B) T-cells showing the sex-dependent variations in the cellular composition upon exposure to differently flavored e-cig aerosols. Representative flow plots (C) and bar graph (D) showing changes in the mean counts of CD4+ and CD8+ T-cells in the lung tissue digest from male and female mice exposed to differently flavored e-cig aerosols as determined using flow cytometry. Data are shown as mean ± SEM (n = 3/sex/group). *p<0.05, and **p<0.01, per Tukey post hoc two-way ANOVA comparison.

Earlier version: Exposure to fruit-flavored e-cig aerosol affected the mitotic pathway genes in the lymphoid cluster

Revised version: Exposure to fruit-flavored e-cig aerosols affect the mitotic pathway genes in the lymphoid cluster

Earlier version: To assess the effects of acute exposures to differently flavored e-cig aerosols on the cell composition in the mouse lungs, we performed differential expression analyses for each flavor in comparison to air and PG:VG controls. The mango-flavored e-cig aerosol exposure had the mildest effect on the cell compositions and gene expression as compared to the controls in our study. We did not observe major dysregulation in the gene expression for myeloid cell cluster in fruit-flavored e-cig exposed mouse lungs as compared to air controls. A total of 24 genes (17 upregulated; 7 downregulated) were differentially expressed in the myeloid clusters in lungs of animals exposed to fruit-flavored e-cig aerosols (Figure 6Ai, Supplementary File S5A). GO analyses identified terms like ‘myeloid cell differentiation’, gas transport’ and ‘H2O2 catabolic process’ as the top hits for this cluster (Figures 6Aii, Supplementary File 5B). Importantly, few gene clusters (Hbb-bs, Hbb-bt, Hba-a2 and Hbb-a1) showed sex-specific changes in the gene expressions in exposed lungs as compared to control. These gene clusters enrich for ‘erythrocyte development’ upon GO analyses and warrant further study. We further found dysregulation of 112 genes (40 up- and 72-down-regulated) in mouse lungs exposed to fruit-flavored e-cig aerosols in the lymphoid cluster. Dysregulation of genes (Incenp, Wnt4, Ccnb2, Trat1, Themis) enriched for ‘T-cell receptor signaling’ and ‘nuclear division’ were observed in exposed group as compared to air control (Figure 6Bi, Supplementary file 9A). Upregulation of genes like Malt1 (Mucosa-associated lymphoid tissue lymphoma translocation protein 1), Ppp1r2 (Protein phosphatase inhibitor 2), and Spib (Transcription factor Spi-B) indicates activation of T-cell mediated immune response upon exposure to fruit-flavored aerosol (24–26). GO analyses showed terms like ‘mitotic nuclear division’, ‘mitotic spindle assembly’, and ‘nuclear chromosome segregation’ to be highly affected on exposure to fruit-flavored e-cig aerosol in C57BL/6J mouse lungs (Figure 6Bii, Supplementary File 9B).

Revised version: To assess the effects of acute exposures to differently flavored e-cig aerosols on the cell composition in the mouse lungs, we performed differential expression analyses for each flavor in comparison to air and PG:VG controls. The fruit-flavored e-cig aerosol exposure had the mildest effect on the cell compositions and gene expression as compared to the controls in our study. We did not observe major dysregulation in the gene expression for myeloid cell cluster in fruit-flavored e-cig exposed mouse lungs as compared to air controls. A total of 24 genes (17 upregulated; 7 downregulated) were significantly (p < 0.05) differentially expressed in the myeloid clusters in the lungs of animals exposed to fruit-flavored e-cig aerosols (Figure 7Ai, Supplementary File S5). GO analyses of differentially expressed (DE) up-regulated genes showed enrichment of terms like ‘NK cell mediated immune response to tumor cells’ (fold enrichment= 9.75; padj = 0.0003), ‘haptoglobin binding’ (fold enrichment= 7.76; padj = 0.0096) and ‘transmembrane-ephrin receptor activity’ (fold enrichment= 7.05; padj = 0.014) in (Figures 7Aii, Supplementary File S6A-B) fruit-flavored aerosol exposed mice lung as compared to control. Importantly, few gene clusters (Hbb-bs, Hbb-bt, Hba-a2 and Hbb-a1) showed sex-specific changes in the gene expressions in exposed lungs as compared to control. These gene clusters enriched for ‘erythrocyte development’ and warrant further study.

Exposure to fruit flavored e-cig aerosols result in activation of oxidative stress-mediated innate immunity in C57BL/6J mouse lungs.

Earlier version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to fruit-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (fruit-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) DEGs in the myeloid (A) and lymphoid (B) cell cluster from fruit-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Revised version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to fruit-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (fruit-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) up/downregulated DEGs in the myeloid (A) and lymphoid (B) cell cluster from fruit-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Added in the revised version: We further found dysregulation of 112 genes (40 up– and 72-down-regulated) in the lymphoid cluster of exposed mice. Upregulation of H2-DMb2, H2-DMb1, Gp5, Pdpn among others enriched for ‘MHC class II protein complex assembly’ (fold enrichment= 7.10; padj = 0.0009) and ‘positive regulation of platelet activation’ (fold enrichment= 6.08; padj = 0.0023) in exposed group as compared to air control. We also observed downregulation of genes like Mcm4, Mcm2, Kif4, Cdca8, Cacna1f, Cacnb3, in the lymphoid cluster of mice exposed to fruit-flavored e-cig aerosol. GO analyses of the downregulated genes enriched for terms like ‘CMG complex’ (fold enrichment= 10.75; padj = 3.39E-07), ‘spindle elongation’ (fold enrichment= 7.72; padj = 0.0005) and ‘high voltage-gated calcium channel activity’ (fold enrichment= 7.09; padj = 0.0059) in fruit-flavored aerosol exposed mouse lungs as compared to air controls(Figure 7Bi-ii, Supplementary File S12-S13).

Earlier version: Exposure to Menthol-flavored e-cig aerosols affected the immune cell function

Revised version: Exposure to Menthol-flavored e-cig aerosols affects the immune cell function

Added in the revised version: We demonstrated an upregulation of 220 genes and a downregulation of 104 genes in the myeloid cluster of mouse lungs exposed to menthol-flavored e-cig aerosol as compared to ambient air. We observed increased expression of inflammatory genes including Il12b, Arid5a, Il12a, Il1b, Cacna1d and Cacnb2 enriching for terms like ‘T-helper 1 cell cytokine production’ (fold enrichment= 6.21; padj = 0.0115) and ‘L-type voltage-gated calcium channel complex’ (fold enrichment= 5.58; padj = 0.0191) in the myeloid cluster of menthol-flavored e-cig aerosol exposed mouse lungs as compared to air. We also found a downregulation in the expression of Ovol1, Mapk15, Erbb2, Nrg4, Katnal2, Hspa1b and Hspa1a in the myeloid cells of flavored e-cig aerosols. GO analyses of the DEGs, thus, showed enrichment of terms like ‘regulation of meiotic cell cycle phase transition’ (fold enrichment= 6.48; padj = 0.0051), ‘ERBB4 signaling pathway’ (fold enrichment= 5.40; padj = 0.025), and ‘protein-containing complex destabilizing activity’ (fold enrichment= 6.36; padj = 0.0095) as the top hits as shown in Figure 8Ai-ii (Supplementary File S7-S8).

Removed in the revised version: We demonstrated an upregulation of 220 genes and a downregulation of 104 genes in the myeloid cluster of mouse lungs exposed to menthol-flavored e-cig aerosol as compared to ambient air. We observed increased expression of inflammatory genes including Cxcl3, Il1r2, Stat4, pointing towards the activation of chemokine-mediated signaling due to exposure to menthol-flavored e-cig aerosol. We also found a significant decrease in the expression of Edn1 (endothelin1) in the myeloid cells of flavored e-cig aerosols as compared to air controls. These genes play a role in positive regulation of neutrophil recruitment and inflammatory responses into the lungs (27–31). GO analyses of the DEGs (p<0.01), thus, showed ‘cytokine-mediated signaling’, ‘cell chemotaxes’, and ‘negative regulation of MAPK cascade’ as the top hits as shown in Figure 7A (Supplementary File 6A-B).

Exposure to Menthol flavored e-cig aerosols result in activation of innate immune responses C57BL/6J mouse lungs.

Earlier version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to menthol-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (menthol-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) DEGs in the myeloid (A) and lymphoid (B) cell cluster from menthol-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Revised version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to menthol-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (menthol-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) up/downregulated DEGs in the myeloid (A) and lymphoid (B) cell cluster from menthol-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Earlier version: Contrary to the responses observed for exposure to fruit-flavored e-cig aerosols, we found significant upregulation in the expression of Cdk8 and Camk1d genes in the lymphoid cell populations for menthol-flavored aerosol exposed mouse lungs (Figure 7B, Supplementary File 10A). Cdk8 (cyclin-dependent kinase 8) is a transcriptional regulator that has a role in the cell cycle progression (32). Whereas Camk1d (calcium/calmodulin dependent protein kinase ID) functions to regulate calcium-mediated granulocyte function and respiratory burst within the cells(33). Taken together, our results point towards an increase in cell proliferation and gene transcription in the lymphoid cluster of mouse lungs exposed to menthol-flavored e-cig aerosols as indicated by enrichment of terms like ‘cyclin-dependent protein serine/threonine kinase activity’, RNA polymerase II CTD modifying activity’ and ‘calmodulin-dependent protein kinase activity’ associated with this cell cluster (Figure 7Bii, Supplementary File 10B).

Revised version: Contrary to the responses observed for exposure to fruit-flavored e-cig aerosols, we found significant (p< 0.05) upregulation in the expression of Cdk8 and Camk1d genes in the lymphoid cell population for menthol-flavored aerosol exposed mouse lungs (Figure 8Bi, Supplementary File S14). Cdk8 (cyclin-dependent kinase 8) is a transcriptional regulator that has a role in the cell cycle progression (28). Whereas Camk1d (calcium/calmodulin dependent protein kinase ID) functions to regulate calcium-mediated granulocyte function and respiratory burst within the cells(29). GO analyses of up and downregulated genes showed enrichment of terms including ‘regulation of toll-like receptor 9 signaling pathway’ (fold enrichment= 5.96; padj = 0.034), ‘transforming growth factor beta activation’ (fold enrichment= 5.42; padj = 0.044), ‘mitotic DNA replication’ (fold enrichment= –7.98.; padj = 2.51E-05) and ‘outer kinetochore’ (fold enrichment= –8.58.; padj = 1.13E-08) in the lymphoid cluster from menthol-flavored e-cig aerosol exposed mouse lungs (Figure 8Bii, Supplementary File S15A-B).

Exposure to Tobacco-flavored e-cig aerosol elicits immune response in myeloid cell and cell cycle arrest in lymphoid cell population

Earlier version: Like menthol-, tobacco-flavored e-cig aerosol also elicited a significant increase in the expression of 338 genes and a decrease in 215 genes as compared to air controls in the myeloid cell cluster. We observed an increase in the expression of chemokines like Stat4, Il1b, and Il1bos in the myeloid cells resulting in terms like ‘cytokine-mediated signaling pathway’ enriched upon functional annotation of DEGs. Like menthol-exposed mouse lungs, we found significant downregulation of Edn1 gene in the myeloid cluster in lungs of mice exposed to tobacco-flavored e-cig aerosol (Figure 8A, Supplementary File 7A-B).

Revised version: Like menthol-, tobacco-flavored e-cig aerosol also elicited a significant increase in the expression of 338 genes and a decrease in 215 genes as compared to air controls in the myeloid cell cluster. We observed an increase in the expression of chemokines like Stat4, Il1b, Il1bos, Il18r1, Unc13d, Lgals9 and Nkg7 in the myeloid cells resulting in enrichment of terms like ‘T-helper 1 cell cytokine production’ (fold enrichment= 6.66; padj = 0.0011) and ‘natural killer cell degranulation’ (fold enrichment= 6.28; padj = 0.0042) in tobacco-flavored e-cig aerosol exposed mouse lungs as compared to air controls.

Added in the revised version: We further found downregulation of genes including Mcmdc2, Rad51, Spp1 and Slc34a2 which enrich for terms like ‘double-strand break repair involved in meiotic recombination’ (fold enrichment= 6.44; padj = 0.0019) and ‘intracellular phosphate ion homeostasis’ (fold enrichment= 6.52; padj = 0.0052) in myeloid cluster from tobacco-exposed e-cig aerosols (Figure 9Ai-ii, Supplementary File S9-S10).

Exposure to Tobacco flavored e-cig aerosols result in activation of cytolysis and neutrophil chemotaxis in C57BL/6J mouse lungs.

Earlier version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to tobacco-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (tobacco-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) DEGs in the myeloid (A) and lymphoid (B) cell cluster from tobacco-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Revised version: Male and female C57BL/6J mice were exposed to 5-day nose-only exposure to tobacco-flavored e-cig aerosols. The mice were sacrificed after the final exposure and mouse lungs from air (control) and aerosol (tobacco-flavored) exposed groups was used to perform scRNA seq. Heatmap and bar plot showing the DESeq2 (i) and GO analyses (ii) results from the significant (p<0.05) up/downregulated DEGs in the myeloid (A) and lymphoid (B) cell cluster from tobacco-flavored e-cig aerosol exposed mouse lungs as compared to controls.

Added in the revised version: We also observed a downregulation of genes responsible for chaperone-mediated protein folding (Cct5, Cct7, Cct8) in the lymphoid cluster from tobacco-flavored e-cig aerosol exposed mouse lungs. Downregulation of these genes could be indicative of the accumulation of misfolded proteins in these lungs which may lead to enhanced cell death (30, 31). DEG and GO analyses identified upregulation of Robo1, Trem2, Padi2, Gp5, Gp9 and Pla2g4a and downregulation of Mcm4, Mcm2, Rad51, Cacnb2, and Cacnb3 in the lymphoid cluster from tobacco-flavored e-cig aerosol exposed mouse lungs resulting in enrichment of terms including ‘regulation of chemokine-mediated signaling pathway’ (fold enrichment= 5.43; padj = 0.0304), ‘positive regulation of platelet activation’(fold enrichment= 5.43; padj = 0.0089), ‘mitotic DNA replication’ (fold enrichment= –9.6; padj = 6.63E-07), and ‘L-type voltage-gated calcium channel complex’ (fold enrichment= –9.12; padj = 0.0005) (Figure 9Bi-ii, Supplementary File S16-S17).

Removed in the revised version: We also observed a downregulation of genes responsible for chaperone-mediated protein folding (Cct5, Cct7, Cct8) in the lymphoid cells from tobacco-flavored e-cig aerosol exposed mouse lungs. Downregulation of these genes could be indicative of the accumulation of misfolded proteins in these lungs which may lead to enhanced cell death (34, 35). In fact, we found increased expression of killer cell lectin-like receptor (Klra)-4 and 8 in exposed mice lungs as compared to air controls, thus indicating that upon exposure to tobacco-flavored e-cig aerosols, the lymphoid cells undergo protein misfolding thereby resulting in increased cell death (Figure 8B, Supplementary File 10A-B).

Dysregulation of chemokine signaling and T-cell activation on exposure to flavored e-cig aerosols

Earlier version: Since we showed increased production of cytokines/chemokines, driving the immune responses in mouse lungs exposed to flavored e-cigs, we performed multianalyte assay to determine the levels of these inflammatory cytokines in the lung digests from the exposed animals as shown in Figure S3D. Exposure to tobacco-flavored e-cig aerosol resulted in a marked increase in the levels of chemotactic chemokines including CXCL16, CXCL12, CXC3R and proinflammatory cytokines including CCl12, CCL17, CCL24, and Eotaxin in the mouse lung digests as compared to air control. Supporting our transcriptional data, the levels of IL1b were dampened in the exposed mouse lungs thus showing negative regulation of IL-1 mediated signaling. We identified sex-dependent changes in the cytokine levels in lung digests from fruit-flavored e-cig aerosols with male mice showing more dampening of cytokine/chemokine protein levels as compared to their female counterparts which aligns with lowed innate immune responses driven by neutrophils and macrophages in the males. Interestingly, the fold changes in the PG:VG+Nic group were contrasting to those observed by PG:VG alone and flavored e-cig aerosol exposed mouse lungs, again supporting our previous deduction that the observed changes are observed as a concerted effect of chemical present in the e-liquid used for flavored e-cig aerosol exposure.

Revised version: Since we showed increased production of cytokines/chemokines, driving the immune responses in mouse lungs exposed to flavored e-cigs, we performed multianalyte assay to determine the levels of these inflammatory cytokines in the lung digests from the exposed animals as shown in Figure S4E. Exposure to tobacco-flavored e-cig aerosol resulted in a marked increase in the levels of chemotactic chemokines including CXCL16, CXCL12, CXC3R and proinflammatory cytokines including CCl12, CCL17, CCL24, and Eotaxin in the mouse lung digests as compared to air control. Interestingly, the fold changes in the PG:VG+Nic group were contrasting to those observed by PG:VG alone and flavored e-cig aerosol exposed mouse lungs, but none of these changes were highly significant.

Earlier version: To identify genes that were commonly dysregulated upon exposure, we generated a list of common genes that were significantly dysregulated in exposure categories (fruit, menthol and tobacco). We identified 9 such target genes-Neurl3, Egfem1, Stap1, Tfec, Mitf, Cirbp, Hist1h1c, Gmds, and Htr2c that were dysregulated in the myeloid cluster from lungs exposed to differently flavored e-cig aerosol, but not PG:VG. We observed significant upregulation of Neurl3, Stap1, Cirbp, and Hist1h1c and downregulation of Tfec, Mitf, Gmds, and Htr2c in the myeloid cluster of mice exposed to differently flavored e-cig aerosols as compared to air control (Supplementary Figure 6A). On analyzing the lymphoid cluster for commonly dysregulated genes, we identified - Klra8 (Killer cell lectin-like receptor 8) and Nfia (nuclear factor I)- that were significantly upregulated in the exposure groups as compared to air-controls (Figure S6A). Klra8 is a natural killer cell associated gene, and its upregulation is generally associated with viral infection associated host immune responses within the mouse lungs (36–38). Nfia is a transcriptional activator responsible for regulating Oxphos-mediated mitochondrial responses and proinflammatory pathways (39, 40).

Revised version: To identify genes that were commonly altered upon exposure, we generated a list of common genes that were significantly dysregulated in exposure categories (fruit, menthol and tobacco). We identified 9 such target genes-Neurl3, Egfem1, Stap1, Tfec, Mitf, Cirbp, Hist1h1c, Gmds, and Htr2c- that were dysregulated in the myeloid cluster from lungs exposed to differently flavored e-cig aerosol, but not PG:VG. We observed significant upregulation of Neurl3, Stap1, Cirbp, and Hist1h1c and downregulation of Tfec, Mitf, Gmds, and Htr2c in the myeloid cluster of mice exposed to differently flavored e-cig aerosols. Upon analyzing the lymphoid cluster for commonly dysregulated genes, we identified – Klra8 (Killer cell lectin-like receptor 8) and Nfia (nuclear factor I)-that were significantly upregulated in the exposure groups as compared to air-controls (Figure S5A). Klra8 is a natural killer cell associated gene, and its upregulation is generally associated with viral infection associated host immune responses within the mouse lungs (32–34). Nfia is a transcriptional activator responsible for regulating Oxphos-mediated mitochondrial responses and proinflammatory pathways (35, 36).

Earlier version: Overall, we identified a total of 29 commonly dysregulated gene targets that were identified from five major cell clusters and performed gene enrichment analyses on the identified targets to identify the top hits (Tables 2-3). Terms like ‘negative regulation of immune system’ (Hmgb3/Gpam/Scgb1a1/Stap1/Ldlr), ‘positive regulation of lipid biosynthetic pathway’ (Htr2c/Gpam/Ldlr), and ‘receptor recycling’ (Ldlr/Ramp3) were amongst the top hits in our observations (Figure S6B).

Revised version: Overall, we identified a total of 29 commonly dysregulated gene targets that were identified from five major cell clusters and performed gene enrichment analyses on the identified targets to identify the top hits (Tables 2-3). Terms like ‘negative regulation of immune system’ (Hmgb3/Gpam/Scgb1a1/Stap1/Ldlr), ‘positive regulation of lipid biosynthetic pathway’ (Htr2c/Gpam/Ldlr), and ‘receptor recycling’ (Ldlr/Ramp3) were amongst the top hits in our observations (Figure S5B).

Earlier version: Of note, the data presented in this study is a sub-part of a larger study. In addition to the groups mentioned in this manuscript, we also had two additional groups of Tobacco-Derived (TDN) and Tobacco-Free Nicotine (TFN). Though further objectives and experimentations performed in both these studies were varied, common air and PG:VG samples were used for analyses of cytokine/chemokine and cell count data as described in our recent publication(41).

Revised version: Of note, the data presented in this study is a sub-part of a larger study. In addition to the groups mentioned in this manuscript, we also had two additional groups of TDN and Tobacco-Free Nicotine (TFN). Though further objectives and experimentations performed in both these studies were varied, common air and PG:VG samples were used for analyses of cytokine/chemokine and cell count data as described in our previous publication(24).

Discussion

Earlier version: E-cigs and associated products have constantly been under scrutiny by the US Food and Drug Administration (FDA) due to public health concerns. In February 2020, FDA placed a regulation on all cartridge-based flavored e-cigs except for menthol and tobacco to reduce the use of e-cigs amongst adolescents and young adults. But it left a loophole for the sale of flavored (including menthol) disposable and open system e-cigs (1, 42). Importantly, most e-cig related bans in the US happened at the state level, thus allowing differential levels of restrictions imposed on the premarket tobacco applications (PMTAs) and sales which defeats the purpose of limiting their accessibility to the general public (43, 44). In fact, the use of nicotine containing e-cigs amongst youth and associated policy restrictions have recently been found to be linked to unintended increase in traditional cigarette use (45). Each year new products are introduced in the market with newer device designs and properties, to lure the users (adults between the ages of 18-24 years) which makes it crucial to continue with the assessments of toxicity and health effects of e-cigs in an unbiased manner (46).

Revised version: E-cigs and associated products have constantly been under scrutiny by the US Food and Drug Administration (FDA) due to public health concerns. In February 2020, FDA placed a regulation on all cartridge-based flavored e-cigs except for menthol and tobacco to reduce the use of e-cigs amongst adolescents and young adults. But it left a loophole for the sale of flavored (including menthol) disposable and open system e-cigs (1, 37). Importantly, most e-cig related bans in the US happened at the state level, thus allowing differential levels of restrictions imposed on the premarket tobacco applications (PMTAs) and sales which defeats the purpose of limiting their accessibility to the general public (38, 39). In fact, the use of nicotine containing e-cigs amongst youth and associated policy restrictions have recently been found to be linked to unintended increase in traditional cigarette use (40). Each year new products are introduced in the market with newer device designs and properties to lure the users (adults between the ages of 18-24 years) which makes it crucial to continue with the assessments of toxicity and health effects of e-cigs in an unbiased manner (41).

Earlier version: Numerous studies indicate increased oxidative stress, DNA damage, and loss of neutrophil function due to exposure to e-cig aerosols in vitro and in vivo (7, 41, 47–50). However, we do not have much knowledge about the cell populations and biological signaling mechanisms that are most affected upon exposure to differently flavored e-cig aerosols at a single-cell level. To bridge this gap in knowledge, we studied the transcriptional changes in the inflammatory responses due to acute (1-hr nose-only exposure including 120 puffs for 5-consecutive days) exposure to fruit-, menthol- and tobacco-flavored e-cig aerosols using single-cell technology.

Revised version: Numerous studies indicate increased oxidative stress, DNA damage, and loss of neutrophil function due to exposure to e-cig aerosols in vitro and in vivo (7, 24, 42–45). However, we do not have much knowledge about the cell populations and biological signaling mechanisms that are most affected upon exposure to differently flavored e-cig aerosols at a single-cell level. To bridge this gap in knowledge, we studied the transcriptional changes in the inflammatory responses due to acute (1-hr nose-only exposure including 120 puffs for 5-consecutive days) exposure to fruit-, menthol– and tobacco-flavored e-cig aerosols using single-cell technology. Short-term (2 hr/day for 3 or 5 consecutive days) e-cigarette exposure studies conducted by others and our group have shown increase in the pro-inflammatory cytokines and oxidative stress markers in both pulmonary and vascular tissues (25, 46, 47). Since nose-only exposures are more direct than whole-body, we reduced the daily exposures of mice to one hour for 5 consecutive days. This dose and duration was found to show sex-dependent changes in MMP-2 and –9 activity and expression (24). Prior literature formed the basis of the current study to assess the effect of acute exposure to differently flavored e-cig aerosols on the lung cell population using a nose-only exposure system.

Earlier version: Reports indicate that the release of metal ions due to the burning of metal coil is a major source of variation during e-cig exposures (51–53). A 2021 study reported the presence of 21 elements in the pod atomizers from different manufacturers identifying a high abundance of 11 elements including nickel, iron, zinc, and sodium amongst others (54). Previous studies have also shown the presence of similar elements in the e-cig aerosols which could have a possible adverse health effect on the vapers (19, 55). More importantly, our study points towards a much important issue, which pertains to the product design of the e-cig vapes. We believe that the aerosol composition varies based on the type of atomizer, coil resistance, coil composition, and chemical reactivity of the e-liquid being used. While much work is done on the chemical composition of the flavors and e-liquids, the other aspects of device design remain understudied and must be an area of research in the future. Another factor that may limit the interpretation of our results pertains to the correlation between the leached metal and the observed transcriptional changes. Since our study provides proof of day-to-day variation in the leaching of metal ions from the same liquid using the same atomizer, it could be possible to develop a statistical model correlating the differential metal exposure to the gene expression changes. We do not conduct such analyses as this was not the focus of this work, but it is a possibility that could be explored in the future.

Revised version: Reports indicate that the release of metal ions due to the burning of metal coil is a major source of variation during e-cig exposures (48–50). A 2021 study reported the presence of 21 elements in the pod atomizers from different manufacturers identifying a high abundance of 11 elements including nickel, iron, zinc, and sodium amongst others (51). Previous studies have also shown the presence of similar elements in the e-cig aerosols which could have a possible adverse health effect on the vapers (19, 52). More importantly, our study points towards a much important issue, which pertains to the product design of the e-cig vapes. We believe that the aerosol composition varies based on the type of atomizer, coil resistance, coil composition, and chemical reactivity of the e-liquid being used. While much work is done on the chemical composition of the flavors and e-liquids, the other aspects of device design remain understudied and must be an area of research in the future. In fact, a recent publication by our group emphasizes this aspect of product design by studying the exposure profile from low and high-resistance coil in an e-cig (21). Another factor that may limit the interpretation of our results pertains to the correlation between the leached metal and the observed transcriptional changes. Since our study provides proof of day-to-day variation in the leaching of metal ions from the same liquid using the same atomizer in the emitted aerosols, it could be possible to develop a statistical model correlating the differential metal exposure to the gene expression changes to get a good assessment of the dose of each toxicant deposited on the lung surface, which has not been done in this study. This raises a pertinent question about what parameters should be controlled to design a comparative study between different flavors of e-cigs in the market per the standards of particle distribution and characteristics. Flavor-dependent variations in daily leaching of metals from an e-cig coil may result in variable deposition of toxicants onto the epithelial lining of both mouth and lungs in frequent vapers. Thus, future studies need to consider conditions like flavor, wattage, coil resistance, coil composition and atomizer life when designing in vitro and in vivo studies to deduce the acute and chronic toxicities of vaping in human.

Earlier version: E-cig vaping has been known to affect the innate and adaptive immune responses among vapers (49, 56–58), but flavor-specific effects on immune function are not fully explored. While we expected to see flavor-dependent changes in our experiment, we did not anticipate observing interesting sex-dependent variation in the lung tissues at single-cell level. In this respect, recent studies show concurring evidence suggesting sex-specific changes in lung inflammation, mitochondrial damage, gene expression, and even DNA methylation in mice exposed to e-cig aerosols (59, 60).

Revised version: E-cig vaping has been known to affect the innate and adaptive immune responses among vapers (44, 53–55), but flavor-specific effects on immune function are not fully explored. While we expected to see flavor-dependent changes in our experiment, we did not anticipate observing interesting sex-dependent variation in the lung tissues using flow cytometry. In this respect, recent studies show concurring evidence suggesting sex-specific changes in lung inflammation, mitochondrial damage, gene expression, and even DNA methylation in mice exposed to e-cig aerosols (56, 57).

Added in the revised version: It is important to highlight here that while the changes in macrophage and neutrophil frequencies ascertained by scRNA seq corroborate with that observed using flow cytometry; similar correlation was not observed for the CD4+ and CD8+ T cell data. The possible explanation for such a discrepancy could be high gene dropout rates in scRNA seq (58), different analytical resolution for the two techniques (59) and pooling of samples in our single cell workflow. Also, while we were able to identify some interesting changes in the eosinophil population within the lungs upon exposure to e-cigs using our flow cytometry data, we could not identify a cluster for eosinophils in our scRNA seq dataset. This could be due to the loss of this cell type during the sample preparation during scRNA seq capture or filtering out as empty droplets during data analyses. Theoretically, eosinophils are known to have lower RNA content and express fewer number of genes, which may result in them being identified as empty droplets and removed when running Cell Ranger (60). Despite the lack of data from scRNA seq, our findings are of relevance as an increased neutrophilic, but a decreased eosinophilic response is characteristic of severe inflammation, infection and asthma (61–63). Further probing into the role of these two cell types in shaping the immune landscape within the lung upon e-cig aerosol exposure is paramount in understanding the specific toxicities and impact of each e-cig flavor on human health and well-being.

Earlier version: One of the most interesting discoveries from our single-cell analyses was the identification of a cluster of immature neutrophils without Ly6G surface marker. While we observed an increase in the cell percentages of these cells in our treatment groups, little to no change in the gene expression was noted which could be indicative of impaired function of these neutrophil population, a fate being reported by various previous studies pertaining to e-cig exposures (13, 56). In fact, our study identified an increase in the level of both mature (Ly6G+) and immature (Ly6G-) neutrophils in tobacco-exposed mouse lung which was corroborated by (a) increase in staining for Ly6G – and Ly6G+ neutrophils in female, but not male mice, through immunofluorescence, and (b) mild to no change in the levels of IL-1β in the lung tissue digests from exposed mice as compared to air. Ly6G is an important marker of neutrophil maturation in mammalian cells and has been reported in relation to various bacterial and parasitic infections in previous studies (61, 62).Exposure to tobacco flavored e-cig aerosols may result in a loss of function of mature neutrophils in the mouse lungs which gets replaced by inactive/immature neutrophils that are deficient in Ly6G surface marker for maturation. In fact, this could explain our inability to identify significant changes in mature neutrophil population through flow cytometry (where Ly6G was used as a marker) when comparing treated and control groups. Interestingly, our in vivo findings were corroborated by a recently published work where neutrophils from healthy volunteers demonstrated a reduction in neutrophil chemotaxis, phagocytic function, and neutrophil extracellular trap formation on exposure to e-cig aerosols, thus suggesting loss of function by mature neutrophils (49). However, this study did not investigate the effects on the immature neutrophil population, an area that needs further study in future research.

Revised version: One of the most interesting discoveries from our single-cell analyses was the identification of a cluster of immature neutrophils without Ly6G surface marker. While we observed an increase in the cell percentages of these cells in our treatment groups, little to no change in the gene expression was noted (data not shown) which could be indicative of impaired function of mature neutrophil population, a fate being reported by various previous studies pertaining to e-cig exposures (13, 53). In fact, our study identified (a) a moderate increase in the neutrophil count in menthol and tobacco-flavored e-cig aerosol treated mouse lungs through scRNA seq analyses, (b) corroboration of the findings from scRNA seq using flow cytometry with more pronounced increase in Ly6G+ neutrophils for male menthol-flavored e-cig aerosol exposed mice and (c) identification of decreased co-localization for Ly6G and S100A8 in the lungs of tobacco-flavored e-cig aerosols as compared to air control through co-immunostaining. Ly6G is an important marker of neutrophil recruitment and maturation in mammalian cells and has been reported in relation to various bacterial and parasitic infections in previous studies (64, 65). Of note, this cluster did not express SiglecF and had high expression of other neutrophil markers including S100A8/9, Lcn2 and Il1b, thus negating the possibility of eosinophils being misrepresented as ‘Ly6G-neutrophils’ in our analyses.

Added in the revised version: S100A8 act as damage associated molecular patterns (DAMPs) and are responsible for neutrophil activation and neutrophil extracellular trap (NET) formation (66, 67). While scRNA seq results in our study identified expression of S100A8/A9 genes in both neutrophil clusters –Ly6G+ and Ly6G-, the expression of Ly6G was totally absent for clusters identified as Ly6G-. Co-immunoprecipitation results also showed expression of both S100A8 and Ly6G markers within the lungs of treated and control lungs, but the co-localization of the two markers was more prominent in the control lungs, thus pointing towards a possible shift in the neutrophil function and activity upon exposure to e-cig aerosols. We are not the first to report the Ly6G deficiency amongst neutrophils.

Added in the revised version: Previous work by Deniset et al (2017) and Kleinholz et al (2021) have highlighted the importance of Ly6G deficiency in relation to infection (64, 65). Deniset et al in their 2017 study reported the presence of two populations of neutrophils in the splenic tissue of Streptococcus pneumoniae infected mice, based on the Ly6G expression– Ly6Ghigh and Ly6Gintermediate. They found that while the former corresponds to mature neutrophils with ability of tissue migration and bacterial clearance, the latter constitutes the immature, immobile pool of neutrophils responsible for proliferation and replenishment of the mature pool of neutrophils (64). The later study by Kleinholz et al (2021) demonstrated decreased uptake of Leishmania major in Ly6G deficient mice thus leading to delayed recruitment to and pathogen capture by neutrophils at the site of infection. Taken together these studies, provide evidence for a protective/compensatory role of the loss of Ly6G in the neutrophil population (65). This in addition to the recent findings by Jasper et al (2025) where neutrophils from healthy volunteers demonstrated a reduction in neutrophil chemotaxis, phagocytic function, and NET formation upon exposure to e-cig aerosols (44), point towards a probable shift in the neutrophil dynamics upon exposure to e-cig aerosols. It is important to mention here that though our results from co-immunostaining using Ly6G and S100A8 pointed towards a shift in innate immune responses upon exposure to e-cig aerosols, it must not be confused with them being only expressed by neutrophil population. S100A8 is expressed in myeloid population including neutrophils, monocytes and macrophages (68, 69) and a subgroup of eosinophils are also known to express Ly6G (70). Thus, an in-depth characterization of these identified population is important to understand the cellular and molecular responses towards e-cig aerosol exposure in vivo.

Earlier version: Of note, myeloid and lymphoid limbs of immunity are interconnected (63, 64). A decline in the neutrophil-mediated immunity might activate other cell types to offer protection. In our study, we find a decline in the neutrophilic immune responses in menthol and tobacco flavored e-cig exposed mouse lungs. Contrarily, we report an increase in the T-cell responses in the form of increased CD8+ T cells from both scRNA seq and flow cytometric analyses in male mice. In fact, increased expression of genes including Malt1, Serpinb9b, and Sema4c are indicative of enhanced T-cell mediated immune response in the lungs of mice exposed to fruit(mango) flavored e-cig aerosols (65–67). Contrary to this, exposure to menthol-flavored e-cig aerosols had a much milder effect on the lymphoid population within the lungs of C57BL/6J mice. We found evidence for increased lymphoid cell proliferation due to activation of cyclin-dependent protein kinase signaling mediated via expression of genes including Cdk8 and Camk1d in these cells (33, 68). Exposure to tobacco-flavored e-cig aerosol provided evidence for decreased chaperone-mediated protein folding, due to the downregulation of Chaperonin Containing TCP-1 (CCT) family of proteins. Chaperonin Containing TCP-1 proteins are important to regulate the production of native active, tubulin, and other proteins crucial for cell cycle progression and cytoskeletal organization (69). This is in conjunction with the upregulation of Klra4 and Klra8 that is indicative of increased protein misfolding and cytotoxic responses in the lymphoid cells of tobacco-exposed e-cig aerosols (70, 71).

Revised version: The myeloid and lymphoid limbs of immunity are interconnected (71, 72). In our study, we find an incidental shift in the neutrophil dynamics in menthol and tobacco flavored e-cig exposed mouse lungs. In contrast, we report an increase in the T-cell responses in the form of increased CD8+ T cells from both scRNA seq and flow cytometric analyses in male mice. In fact, increased expression of genes including Malt1, Serpinb9b, and Sema4c are indicative of enhanced T-cell mediated immune response in the lungs of mice exposed to fruit (mango) flavored e-cig aerosols (73–75). Contrary to this, exposure to menthol-flavored e-cig aerosols had a much milder effect on the lymphoid population within the lungs of C57BL/6J mice. We found evidence for increased lymphoid cell proliferation due to activation of cyclin-dependent protein kinase signaling mediated via expression of genes including Cdk8 and Camk1d in these cells (29, 76). Exposure to tobacco-flavored e-cig aerosol provided evidence for decreased chaperone-mediated protein folding, due to the downregulation of Chaperonin Containing TCP-1 (CCT) family of proteins. Chaperonin Containing TCP-1 proteins are important to regulate the production of native actin, tubulin, and other proteins crucial for cell cycle progression and cytoskeletal organization (77). This is in conjunction with the upregulation of Klra4 and Klra8 that is indicative of increased protein misfolding and cytotoxic responses in the lymphoid cells of tobacco-exposed e-cig aerosols (78, 79).

Earlier version: Overall, we provide evidence of subdued innate immune responses due to loss of function of neutrophils and increased T-cell proliferation and cytotoxicity in a flavor-dependent manner upon exposure to e-cig aerosols in this study. An increase in the levels of CCL17, CCL20, CCL22, IL2 and Eotaxin in the lung digests from tobacco-exposed mouse lungs further support this deduction as these cytokines/chemokines are associated with T-cell mediated immune responses (72–79). Importantly CXCL16 attracts T-cells and natural killer cells to activate cell death. It is involved in LPS- mediated acute lung injury (ALI), an outcome which has also been linked with e-cig exposures in human (80, 81).

Revised version: Overall, we provide evidence of altered innate immune responses due to variable neutrophilic-eosinophilic function and increased T-cell proliferation and cytotoxicity in a flavor-dependent manner upon exposure to e-cig aerosols in this study. An increase in the levels of CCL17, CCL20, CCL22, IL2 and Eotaxin in the lung digests from tobacco-exposed mouse lungs further support this deduction as these cytokines/chemokines are associated with T-cell mediated immune responses (80–87). Importantly CXCL16 attracts T-cells and natural killer cells to activate cell death. It is involved in LPS-mediated acute lung injury (ALI), an outcome which has also been linked with e-cig exposures in human (88, 89).

Earlier version: Further, we compiled a list of commonly dysregulated genes in a flavor independent manner and identified 29 gene targets. Signal-transducing adaptor protein-2 (Stap1) which is commonly upregulated upon exposure to e-cig aerosols is known to regulate T-cell activation and airway inflammation which is in agreement with the overall outcome of our findings (82). Another gene that was found to be consistently dysregulated in many cell types was Cold-inducible RNA binding protein (Cirbp). Cirbp is a stress response protein linked with stressors like hypoxia. Its upregulation upon e-cig exposure supports that vaping induces oxidative stress and can have adverse implications on the exposed cell types (83). 5-hydroxytryptamine receptor 2C (Htr2c) and Klra8 are other genes in this category of commonly dysregulated genes that are associated with enhancing inflammation and cell death (70, 71, 84). Overall, we provide a cell-specific atlas of immune responses upon exposure to differently flavored e-cig aerosols.

Revised version: Further, we compiled a list of commonly dysregulated genes in a flavor independent manner and identified 29 gene targets. Signal-transducing adaptor protein-2 (Stap1) which is commonly upregulated upon exposure to e-cig aerosols is known to regulate T-cell activation and airway inflammation which is in agreement with the overall outcome of our findings (90). Another gene that was found to be consistently dysregulated in many cell types was Cold-inducible RNA binding protein (Cirbp). Cirbp is a stress response protein linked with stressors like hypoxia. Its upregulation upon e-cig exposure supports that vaping induces oxidative stress and can have adverse implications on the exposed cell types (91). Importantly, this gene is involved in DNA repair mechanism thus making it crucial for cell survival pathways (92). 5-hydroxytryptamine receptor 2C (Htr2c) and Klra8 are other genes in this category of commonly dysregulated genes that are associated with enhancing inflammation and cell death (78, 79, 93). Overall, we provide a cell-specific resource of immune responses upon exposure to differently flavored e-cig aerosols.

Earlier version: Considering that scRNA technology has not been commonly used for e-cig research, ours is one of the first studies employing this technique to identify possible changes in the cellular composition and gene expressions. Importantly we use the nose-only exposure system for our experiment to avoid exposure through other routes. However, despite the novel approach and state-of-the-art exposure system, we had a few limitations. First, we used a small sample size to identify the changes in the mouse lungs upon exposure to e-cig aerosols at a single-cell level. Due to the expensive nature of single-cell sequencing technology and limited information in the literature reporting changes at single-cell level, we chose to design this experiment with small sizes of experimental and validation cohort. But, based on the encouraging findings from this study, future studies could be designed with a larger sample size,longer durations of exposure and more targeted approach to identify the acute and chronic effects of vaping in vivo. Second, we could not expand upon the sex-dependent changes observed through our work upon exposure. This was because such an effect was not anticipated when we conceived the idea of a short-term exposure in mice. However, considering the evidence from the cuttenr study, future experimental designs in our lab will be designed taking sex as a crucial confounder for studying the effects of e-cig exposure in translational contexts. Third, the inclusion of PG:VG + Nic group was streamlined in this study, but in future work including this group for scRNA seq analyses might be crucial to delineate the effects of nicotine alone on gene transcription. Fourth, we did not anticipate changes in the metal release on consecutive days of exposure at the start of our study. Later, our data pointed towards the importance of device design in e-cig exposures. Future studies need to identify the factors that may affect the daily composition of e-cig aerosols and devise a method of better monitoring these possible confounders. However, in this regard our experiment does mimic the real-life scenario, as such variations due to prolonged storage of e-liquid and differences arising due to vape design must be common amongst human vapers.

Revised version: Considering that scRNA technology has not been commonly used for e-cig research, ours is one of the first studies employing this technique to identify possible changes in the cellular composition and gene expressions. Importantly we use the nose-only exposure system for our experiment to avoid exposure through other routes. However, despite the novel approach and state-of-the-art exposure system, we had a few limitations. First, we used a small sample size to identify the changes in the mouse lungs upon exposure to e-cig aerosols at a single-cell level. Due to the expensive nature of single cell sequencing technology and limited information in literature, we chose to design this experiment with small sizes of experimental and validation cohort. But, based on the encouraging findings from this study, future studies could be designed with a larger sample size, longer durations of exposure and more targeted approach to identify the acute and chronic effects of vaping in vivo. Second, we could not expand upon the sex-dependent changes observed through our work upon exposure. This was because such an effect was not anticipated when we conceived the idea of a short-term exposure in mice. However, considering the evidence from the current study, future experimental designs in our lab are considering sex as a crucial confounder for studying the effects of e-cig exposure in translational contexts. Third, the inclusion of PG:VG + Nic group was streamlined in this study, but in future work inclusion of this group for scRNA seq analyses to delineate the effects of nicotine alone on gene transcription is necessary. Fourth, we did not anticipate changes in the metal release on consecutive days of exposure at the start of our study. Later, our data pointed towards the importance of device design in e-cig exposures. Future studies need to identify the factors that may affect the daily composition of e-cig aerosols and devise a method of better monitoring these possible confounders. However, in this regard our experiment does mimic the real-life scenario, as such variations due to prolonged storage of e-liquid and differences arising due to vape design must be common amongst human vapers.

In conclusion, we identified cell-specific changes in the gene expressions upon exposure to e-cig aerosols using single cell technology. We identified a set of top 29 dysregulated genes that could be studied as markers of toxicity/immune dysfunction in e-cig research. Future work with larger sample sizes and sex-distribution is warranted to understand the health impacts of long-term use of these novel products in humans.

Methods

Animals Ethics Statement

All experiments were conducted per the guidelines set by the University Committee on Animal Resources at the University of Rochester Medical Center (URMC). Care was taken to implement an unbiased and robust approach during the experimental design and conduction of each experiment to ensure data reproducibility per the National Institutes of Health standards.

Animals

We ordered 5-week-old pups of male and female C57BL/6J mice from Jackson Laboratory to conduct this experiment. Prior to the start of the experiment, mice were housed at the URMC Vivarium for acclimatization. Thereafter the animals were moved to the mouse Inhalation Facility at URMC for training and exposures.

Earlier version: One week prior to the start of the exposures, mice underwent a five-day nose-only training to adapt themselves to the mesh restraints of the exposure tower. The mouse restraint durations were increased gradually to minimize the animal’s stress and discomfort. Of note, the mouse sacrifice was performed within 8-12 weeks’ age for each mouse group to ensure that the mouse age corresponds to the age of adolescents (12- 17 years) in humans (85, 86). Age and sex-matched animals (n = 2/sex/group) used to perform single cell RNA sequencing (scRNA seq) were considered as the ‘experimental cohort’; whereas another group of age and sex-matched (n = 3/sex/group) mice exposed to air and flavored e-cig aerosol served as ‘validation cohort’ for this study.

Revised version: One week prior to the start of the exposures, mice underwent a five-day nose-only training to adapt themselves to the mesh restraints of the exposure tower. The mouse restraint durations were increased gradually to minimize the animal’s stress and discomfort. Of note, the mouse sacrifice was performed within 8-12 weeks’ age for each mouse group to ensure that the mouse age corresponds to the age of adolescents (12-17 years) in humans (94, 95). Age and sex-matched animals (n = 2/sex/group) used to perform single cell RNA sequencing (scRNA seq) were considered as the ‘experimental cohort’; whereas another group of age and sex-matched (n = 3/sex/group) mice exposed to air and flavored e-cig aerosol served as ‘validation cohort’ for this study.

E-cigarette Device and E-liquid

Earlier version: We utilized an eVic-VTC mini and CUBIS pro atomizer (SCIREQ, Montreal, Canada) with a BF SS316 1.0-ohm coil from Joyetech (Joyetech, Shenzhen, China) for vaping and the inExpose nose-only inhalation system from SCIREQ (SCIREQ, Montreal, Canada) for mouse exposures. Both air and PG:VG exposed mice groups were considered as controls for this experiment. We used commercially available propylene glycol (PG; EC Blend) and vegetable glycerin (VG; EC Blend) in equal volumes to prepare a 50:50 solution of PG:VG. For flavored product exposures, mice were exposed to three different e-liquids – a menthol flavor “Menthol-Mint”, a fruit flavor “Mango” and a tobacco flavor “Cuban Blend”. Of note, all the e-liquids were commercially manufactured with 50mg/mL of tobacco derived nicotine (TDN). So, all treatments have nicotine in addition to the flavoring mentioned respectively. Additionally, we used a mixture of PG:VG with 50mg/mL of TDN as a control for limited experiments to study the effect of nicotine alone in our treatment. This group is labeled as PG:VG+Nic for the rest of the manuscript.

Revised version: We utilized an eVic-VTC mini and CUBIS pro atomizer (SCIREQ, Montreal, Canada) with a BF SS316 1.0-ohm coil from Joyetech (Joyetech, Shenzhen, China) for vaping and the inExpose nose-only inhalation system from SCIREQ (SCIREQ, Montreal, Canada) for mouse exposures. Both air and PG:VG exposed mice groups were considered as controls for this experiment. We used commercially available propylene glycol (PG; EC Blend) and vegetable glycerin (VG; EC Blend) in equal volumes to prepare a 50:50 solution of PG:VG. For flavored product exposures, mice were exposed to three different e-liquids – a menthol flavor “Menthol-Mint”, a fruit flavor “Mango” and a tobacco flavor “Cuban Blend”. Of note, all the e-liquids were commercially manufactured with 50mg/mL of tobacco derived nicotine (TDN). So, all treatments have nicotine in addition to the flavoring mentioned respectively. Additionally, we used a mixture of PG:VG with 50mg/mL of TDN as a control for limited experiments to study the effect of nicotine alone in our treatment. This group is labeled PG:VG+Nic for the rest of the manuscript.

E-cigarette Exposure

Earlier version: Scireq Flexiware software with the InExpose Inhalation system was used for controlling the Joyetech eVic-VTC mini device to perform nose-only mouse exposures. For this exposure, we utilized a puffing profile that mimicked the puffing topography of e-cig users in two puffs per minute with a puff volume of 51 mL, puff duration of 3 seconds, and an inter puff interval of 27 seconds with a 2 L/min bias flow between puffs (87). Figures 1A & S1A depict the experimental design and exposure system employed for this study.

Revised version: Scireq Flexiware software with the InExpose Inhalation system was used for controlling the Joyetech eVic-VTC mini device to perform nose-only mouse exposures. For this exposure, we utilized a puffing profile that mimicked the puffing topography of e-cig users in two puffs per minute with a puff volume of 51 mL, puff duration of 3 seconds, and an inter puff interval of 27 seconds with a 2 L/min bias flow between puffs (96). Figures 1A & S1A depict the experimental design and exposure system employed for this study.

Removed in the revised version: Age-matched male and female (n=5 per sex) mice were used for each group, namely, air, PG:VG, Fruit, Menthol, and Tobacco. To ensure rigor and reproducibility in our work, we have used age- and sex-matched control and treated mice in this study. Confounders like environment and stress were minimized by housing all the cages in environmentally controlled conditions and training all the mice (both control and treated) in nose-only chambers. Each group of mice was exposed to the above-mentioned puffing profile for one hour each day (120 puffs) for a total of five consecutive days.

Earlier version: Additionally, a group of mice (n=3/sex) was exposed to PG:VG + Nic for the same duration using similar exposure profile to serve as control to assess the effect of nicotine on the observed changes using selected experiments. Air-exposed mice were exposed to the same puffing profile for a total of five consecutive days to ambient air. We recorded the temperature, humidity, and CO levels of the aerosols generated at the start, mid, and end of the exposure on each day using the Q-Trak Indoor Air Quality Monitor 7575 (TSI, Shoreview, MN). Total Particulate Matter (TPM) sampling was done from the exhaust tubing of the set-up at the 30 min mark of the exposure and at the inlet connected to the nose-only tower (shown in Figure S1A) immediately after the culmination of the exposure. Gravimetric measurements for TPM were also conducted to confirm relative dosage to each mouse group daily.

Revised version: Age-matched male and female (n=5 per sex) mice were used for each group, namely, air, PG:VG, Fruit, Menthol, and Tobacco. To ensure rigor and reproducibility in our work, we have used age– and sex-matched control and treated mice in this study. Confounders like environment and stress were minimized by housing all the cages in environmentally controlled conditions and training all the mice (both control and treated) in nose-only chambers. Each group of mice was exposed to the above-mentioned puffing profile for one hour each day (120 puffs) for a total of five consecutive days. Additionally, a group of mice (n=3/sex) was exposed to PG:VG + Nic for the same duration using similar exposure profile to serve as control to assess the effect of nicotine on the observed changes using selected experiments. Air-exposed mice were exposed to the same puffing profile for a total of five consecutive days to ambient air. We recorded the temperature, humidity, and CO levels of the aerosols generated at the start, mid, and end of the exposure on each day using the Q-Trak Indoor Air Quality Monitor 7575 (TSI, Shoreview, MN). Total Particulate Matter (TPM) sampling was done from the exhaust tubing of the set-up at the 30 min mark of the exposure and at the inlet connected to the nose-only tower (shown in Figure S1A) immediately after the culmination of the exposure. Gravimetric measurements for TPM were also conducted to confirm relative dosage to each mouse group daily.

Preparation of single-cell suspension

Earlier version: The animals were sacrificed (at 8-10 weeks’ age) immediately after the final exposure. Vascular lung perfusion was performed using 3 mL of saline before harvesting the lung lobes for preparation of single-cell suspension. It is important to mention here that of the 5 lung samples/sex/group; 2 sex/group were used for histological assessments and scRNA analyses. Here, the left lung lobe was inflated using low-melting agarose and used for histology, while the rest of the uninflated lung lobes were used for preparing the single-cell suspension. We pooled the lung lobes for each sex per group for preparation of the single-cell suspension as depicted in Figure 1A. The lung lobes were weighed digested using Liberase method as described earlier (88). Briefly, lung lobes were weighed and digested using Liberase (Cat# 5401127001; Roche, Basel, Switzerland) enzymatic cocktail with 1% DNase. The tubes were then transferred to the gentleMACS dissociator (Miltenyi, Gaithersburg, MD) and the manufacturer’s protocol for mouse lung digestion was run. The sample tubes were next incubated at 370C for 30 minutes with constant rotation after which the suspension was strained through 70- micron MACS Smart Strainer. Thereafter, the suspension was centrifuged at 500g for 10 minutes at 40C, the supernatant was discarded, and 0.5 mL of RBC lysis buffer was added to the cell pellet to digest RBCs. The suspension was left on ice for 5 min in RBC lysis buffer and then 4 mL of ice-cold PBS with 10% FBS was added to stop the lysis. The suspension was again centrifuged at 500g for 10 minutes at 40C. The cell pellet was suspended in 1 mL PBS with 10% FBS, and cell number and viability were checked using AO/PI staining on a Nexcelom Cellometer Auto2000.

Revised version: The animals (at 8-10 weeks’ age) were sacrificed immediately after the final exposure. Vascular lung perfusion was performed using 3 mL of saline before harvesting the lung lobes for preparation of single-cell suspension. It is important to mention here that of the 5 lung samples/sex/group; 2 sex/group were used for histological assessments and scRNA analyses. Here, the left lung lobe was inflated using low-melting agarose and used for histology, while the rest of the uninflated lung lobes were used for preparing the single-cell suspension. We pooled the lung lobes for each sex per group for preparation of the single-cell suspension as depicted in Figure 1A. The lung lobes were weighed and digested using Liberase method as described earlier (97). Briefly, lung lobes were weighed and digested using Liberase (Cat# 5401127001; Roche, Basel, Switzerland) enzymatic cocktail with 1% DNase. The tubes were then transferred to the gentleMACS dissociator (Miltenyi, Gaithersburg, MD) and the manufacturer’s protocol for mouse lung digestion was run. The sample tubes were next incubated at 370C for 30 minutes with constant rotation after which the suspension was strained through 70-micron MACS Smart Strainer. Thereafter, the suspension was centrifuged at 500g for 10 minutes at 40C, the supernatant was discarded, and 0.5 mL of RBC lysis buffer was added to the cell pellet to digest RBCs. The suspension was left on ice for 5 min in RBC lysis buffer and then 4 mL of ice-cold PBS with 10% FBS was added to stop the lysis. The suspension was again centrifuged at 500g for 10 minutes at 4°C. The cell pellet was suspended in 1 mL PBS with 10% FBS, and cell number and viability were checked using AO/PI staining on a Nexcelom Cellometer Auto2000.

Library Preparation and Single-Cell Sequencing

The prepared single-cell suspension was sent to the Genomics Research Center (GRC) at URMC for library preparation and single-cell sequencing. Library preparation was performed from control and treatment groups using the 10X single cell sequencing pipeline by 10X Genomics and 10,000 cells were captured per sample using the Chromium platform. The prepared library was sequenced on NovaSeq 6000 (Illumina, San Diego, CA) at a mean sequence depth of 30,000 reads per cell. Read alignment was performed to GRCm38 Sequence.

Data Analyses

Earlier version: We used the standard Seurat v4.2 analyses pipeline to analyze our data (89). In brief the low-quality cells and potential doublets were excluded from the dataset to create the analyses dataset. The residual features due to the presence of RBCs were corrected before integration. “scTransform” function was used for integration of all the datasets after which the standard Seurat pipeline was used for data normalization of integrated data. “FindVariableGenes” gene function was used to identify the variable genes for dimensionality reduction using PCA function. UMAP was used for dimensionality reduction and clustering of cells.

Revised version: We used the standard Seurat v4.3 analyses pipeline to analyze our data (98). In brief the low-quality cells and potential doublets were excluded from the dataset to create the analyses dataset. The residual features due to the presence of RBCs were corrected before integration. “scTransform” function was used for integration of all the datasets after which the standard Seurat pipeline was used for data normalization of integrated data. “FindVariableGenes” gene function was used to identify the variable genes for dimensionality reduction using PCA function. UMAP was used for dimensionality reduction and clustering of cells.

Earlier version: To identify the unique features and cell clusters within each cell subtype, we used the sub-setting feature within Seurat. After identifying the 5 major cell populations (epithelial, endothelial, stromal, myeloid and lymphoid) in our data sets, each of these cell types were sub-clustered using “subset” function, normalized and re-clustered. Cell annotation for each of the subsets was performed with the help of Tabula Muris database (90, 91). However, some clusters were annotated manually with the help of a literature search as discussed in the results section.

Revised version: To identify the unique features and cell clusters within each cell subtype, we used the sub-setting feature within Seurat. After identifying the 5 major cell populations (epithelial, endothelial, stromal, myeloid and lymphoid) in our data sets, each of these cell types were sub-clustered using “subset” function, normalized and re-clustered. Cell annotation for each of the subsets was performed with the help of Tabula Muris database (99, 100). However, some clusters were annotated manually with the help of a literature search as discussed in the results section.

Earlier version: DESeq2 (V. 1.42.1) was used to perform pseudobulk analyses to identify differentially expressed genes within each group. The ClusterProfiler R package (V. 4.10.1) (92) was employed to perform gene enrichment analyses of the differentially expressed genes.

Revised version: DESeq2 (V.1.42.1) was used to perform pseudobulk analyses to identify differentially expressed genes within each group. Here, genes showing a fold change > 0.5 & < –0.5 along with a padj value < 0.05 were considered significantly dysregulated and plotted as heatmap using GraphPad The ClusterProfiler R package (V. 4.10.1) (101) was employed to perform gene enrichment analyses of the differentially expressed genes (fold change> 0.5 & < –0.5).

Cytokine/chemokine assessment

We used multiplex assay to determine the levels of cytokine/chemokine in the lung homogenates from control and e-cig aerosol exposed mouse lungs using commercially available Bio-Plex Pro Mouse Chemokine Assay (Cat#12009159, Bio-RAD, Hercules, CA) per the manufacturer’s instructions. Approximately 40 mg of mouse lung lobes were homogenized in 300uL of 1X RIPA buffer with 0.1% protease and phosphatase inhibitor. The lung homogenate was stored on ice for 30 min. Following incubation, the homogenate was centrifuged at 15000 rpm for 15 min at 4 degrees Celsius. The supernatant was collected and used for performing the multianalyte assay for determination of cytokine/chemokine levels using Luminex FlexMap3D system. A heatmap after normalization of the measured cytokine/chemokine to the protein amount loaded was plotted.

Lung Histology

The left lung lobe of mice used for scRNA seq were inflated with 1% low melting agarose and fixed with 4% neutral buffered PFA. Fixed lungs were dehydrated, before being paraffin-embedded and sectioned (5μm). Hematoxylin and eosin (H&E) staining was performed by the Histology, Biochemistry, and Molecular Imaging Core at URMC. The H&E stain was observed at 10X magnification using Nikon Elipse-Ni fluorescence microscope. Ten to fifteen random images were captured per sample.

Flow cytometry

Earlier version: Flow cytometry was performed on the cells collected from lung homogenates from air and flavored e-cig aerosol exposed mouse lungs. For analyses of immune cell population in the lung, the lung lobes were digested as described earlier (88). The single-cell suspension thus prepared was used to run flow cytometry using the BD LSRFortessa cell analyzer. Cells were blocked with CD16/32 (Tonbo biosciences 70- 0161-u500, 1:10) to prevent nonspecific binding and stained with a master mix of Siglec F (BD OptiBuild Cat#740280, 1:200), CD11b (Biolegend Cat #101243, 1:200), Ly6G (BD Horizon Cat# 562700, 1:200), CD45 (Biolegend Cat#103126, 1:200), CD11c (Biolegend Cat #117318, 1:200), CD4 (Biolegend Cat#116012, 1:200), and CD8 (eBiosciences Cat#17-0081-82, 1:200). 7AAD (eBiosciences Cat#00-6993-50, 1:10) was used as the nucleic acid dye to detect live and dead cells.

Revised version: Flow cytometry was performed on the cells collected from lung homogenates from air and flavored e-cig aerosol exposed mouse lungs. For analyses of immune cell population in the lung, the lung lobes were digested as described earlier (97). The single-cell suspension thus prepared was used to run flow cytometry using the BD LSRFortessa cell analyzer. Cells were blocked with CD16/32 (Tonbo biosciences 70-0161-u500, 1:10) to prevent nonspecific binding and stained with a master mix of Siglec F (BD OptiBuild Cat#740280, 1:200), CD11b (Biolegend Cat #101243, 1:200), Ly6G (BD Horizon Cat# 562700, 1:200), CD45 (Biolegend Cat#103126, 1:200), CD11c (Biolegend Cat #117318, 1:200), CD4 (Biolegend Cat#116012, 1:200), and CD8 (eBiosciences Cat#17-0081-82, 1:200). 7AAD (eBiosciences Cat#00-6993-50, 1:10) was used as the nucleic acid dye to detect live and dead cells. The gating strategy used for this assay has been depicted in Figure S6.

Metal analyses

To understand the levels of metals released during subsequent days of exposure, we performed Inductively coupled plasma mass spectrometry (ICP-MS) on the e-cig aerosol condensates collected from each day of exposure using a Perkin Elmer ICP-MS model 2000C. The samples were run using a Total Quant KED protocol with 4mL/min Helium flow and externally calibrated using a blank and a 100ppb standard for the 51 elements. The samples were submitted to the Element Analyses facility at URMC, and levels of metals thus detected were plotted.

Ly6G /S100A8 Double staining

Earlier version: To determine the presence of Ly6G+ and Ly6G- neutrophils, FFPE tissue sections from air and tobacco-flavored e-cig aerosol exposed lungs were stained with Ly6G and S100A8. In brief, 2-3 tissue sections per sample were deparaffinized using serial incubation in xylene followed by graded alcohol. Slides were incubated in 1X Citrate Buffer (Cat# S1699, Agilent, Santa Clara, CA) for 10 min at 95°C for antigen retrieval which was followed by incubation at room temperature for 30 min. The slides were next washed with water and permeabilized using a permeabilization buffer (0.1% Triton-X in 1X TBST) for 10 min. Next, the slides were again washed with 1X TBST and blocked using Blocking buffer (5% goat serum in 1X TBST) for 30 min at room temperature. The blocked slides were incubated overnight at 4°C with Ly6G (Cat# 16-9668-85, Invitrogen, dilution: 1:100) and S100A8 (Cat# 26992-1-AP, Proteintech, dilution: 1:200). Next day, the slides were washed with 1X TBST and incubated for 2 hrs. at room temperature with goat anti-rabbit Alexa Fluor 594 (Cat # A11012, Invitrogen) and donkey anti-mouse Alexa Fluor 488 (Cat # A21202, Invitrogen) secondary antibody at 1:1000 dilution. Thereafter the slides were washed and mounted with ProLong Diamond Antifade Mountant with DAPI (Cat# P36962, Invitrogen, Waltham, MA). 6-10 images were captured using Cytation 5 Imaging software (Thermo Fisher, Waltham, MA) at 10X magnification.

Revised version: To determine the various populations of neutrophils in exposed and control groups, FFPE tissue sections from air and tobacco-flavored e-cig aerosol exposed lungs were stained with Ly6G and S100A8. In brief, 2-3 tissue sections per sample were deparaffinized using serial incubation in xylene followed by graded alcohol. Slides were incubated in 1X Citrate Buffer (Cat# S1699, Agilent, Santa Clara, CA) for 10 min at 95°C for antigen retrieval which was followed by incubation at room temperature for 30 min. The slides were next washed with water and permeabilized using a permeabilization buffer (0.1% Triton-X in 1X TBST) for 10 min. Next, the slides were again washed with 1X TBST and blocked using Blocking buffer (5% goat serum in 1X TBST) for 30 min at room temperature. The blocked slides were incubated overnight at 4°C with Ly6G (Cat# 16-9668-85, Invitrogen, dilution: 1:100) and S100A8 (Cat# 26992-1-AP, Proteintech, dilution: 1:200). Next day, the slides were washed with 1X TBST and incubated for 2 hrs. at room temperature with goat anti-rabbit Alexa Fluor 594 (Cat # A11012, Invitrogen) and donkey anti-mouse Alexa Fluor 488 (Cat # A21202, Invitrogen) secondary antibody at 1:1000 dilution. Thereafter the slides were washed and mounted with ProLong Diamond Antifade Mountant with DAPI (Cat# P36962, Invitrogen, Waltham, MA). 6-10 images were captured using Cytation 5 Imaging software (Thermo Fisher, Waltham, MA) at 10X magnification. The ImageJ deconvolution was used for quantifying the fluorescence in green (Ly6G+) and red (S100A8+) channels relative to DAPI (blue channel) and the relative fluorescence was plotted.

Statistical significance

We used GraphPad Prism 10.5.0 for all statistical calculations. All the data plotted in this paper are expressed as mean ± SEM. Pairwise comparisons were done using unpaired t test while one-way analysis of variance (ANOVA) with ad-hoc Tukey’s test was employed for multi-group comparisons. To identify sex-based variations in our treatment groups, Tukey post hoc two-way ANOVA was employed.

Added in the revised version: Code Availability

Added in the revised version: Data collection was performed with mkfastq pipeline in Cell Ranger’s (v7.0.1). Cell Ranger (v7.0.1) was used for cell and gene counting using the default settings. Single-cell analysis was performed using the Seurat R package (v4.3.0) using the recommended workflow.